```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

# E7: Introduction to Computer Programming for Scientists and Engineers

## Lab Assignment 2

For each question, you will have to fill in one or more Python functions. We provide an autograder with a number of test cases that you can use to test your function. Note that the fact that your function works for all test cases thus provided does necessarily guarantee
that it will work for all possible test cases relevant to the question. It is your responsibility
to test your function thoroughly, to ensure that it will also work in situations not covered
by the test cases provided

```python
# Please run this cell, and do not modify the contets
import math
import numpy as np
import matplotlib.pyplot as plt
#plt.style.use('fivethirtyeight')
#plt.rcParams["figure.figsize"] = (20,7)
%matplotlib inline

import hashlib
def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

## Question 1

Write a Python function that calculates the roots of a quadratic equation $ax^2+bx+
c = 0$ of unknown $x$, using the quadratic formula below. 

$$x = \frac{-b\pm \sqrt{b^2-4ac}}{2a}$$

The function output should be a two element tuple. The first element, `elem_1` ,should be the root computed by setting the $\pm$ to $+$, and the second element, `elem_2` ,should be the root computed by setting the $\pm$ to $-$. You may assume $b2 - 4ac \ge 0$.

```python
def solve_quadratic(a, b, c):
    """
    Arguments
    ---------
    a : coefficient of x**2
    b : coefficient of x
    c : constant
    
    Returns
    -------
    elem_1 : larger root of quadratic equation
    elem_2 : smaller root of quadratic equation
    """
    discriminant = ...
    elem_1 = ...
    elem_2 = ...
    return elem_1, elem_2 
```

```python
solve_quadratic(6,-17,12)
```

```python
grader.check("q1")
```

## Question 2.0

Carbon dating is a method for determining the age of an object containing the radioisotope carbon-14. Since carbon-14 is subject to radioactive decay, the number of carbon-14 atoms in a sample will decay over time according to the following exponential decay law:

$$ N(t) = N_0 e^{-\lambda t} $$

where $N_0$ is the number of atoms of the isotope in the original sample, $N(t)$ is the remaining number of atoms after a time $t$, $\lambda$ is the expoential decay constant and $t$ is the time.

Write a function that calculates the remaining fraction of carbon-14 in a sample after a certain amount of time (in years). The remaining fraction is defined as the ratio of the number of atoms at the present time ($N(t)$) to the initial number of atoms ($N_0$). You can re-arrange the decay formula above to find an expression for this ratio.

Assume $\lambda = 0.00012097 \: \text{years}^{-1}$.

```python
def c14_dating(time):
    ratio = ...
    return ratio
```

```python
c14_dating(10000)
```

```python
grader.check("q2.0")
```

## Question 2.1

Now create a more general version of the function in Q2.0 and name the new function `exp_decay`. This new function will perform the same calculation as `c14_dating` but include a default argument for $\lambda$, the expoential decay constant. Call the argument for $\lambda$ `decay_constant`. Set the default value for `decay_constant` equal to the exponential decay constant for c-14.

Write your function in the cell below (that contains the ellipsis `...`).

Your function should be in the form:
``` PYTHON
def exp_decay(...):
    ratio = ...
    return ratio
```

```python
...
```

```python
exp_decay(10000)
```

```python
grader.check("q2.1")
```

## Question 3.0

Write a Python function that computes the Euclidean distances between two sets of points. The x-coordinate for the first point will be stored in a column vector `X_1`, the y-coordinate for the first point in `Y_1`, the x-coordinate for the second point will be stored in a column vector `X_2`, the y-coordinate for the second point in `Y_2`. Note that, for a single pair of points, $(x_1, y_1)$ and $(x_2, y_2)$, the
Euclidean distance can be obtained from the distance formula:

$$ \text{Euclidean distance} = \sqrt{(x_1-x_2)^2 + (y_1-y_2)^2} $$

```python
def euclidean_distance(x1, y1, x2, y2):
    """
    Arguments
    ---------
    x1 : The x-coordinates for the first set of points
    y1 : The y-coordinates for the first set of points
    x2 : The x-coordinates for the second set of points
    y2 : The y-coordinates for the second set of points
    
    Returns
    -------
    distance : Euclidean distance between points
    """
    x = ...
    y = ...
    distance = ...
    return distance
```

```python
euclidean_distance(4,5,3,7)
```

```python
grader.check("q3.0")
```

## Question 3.1

Now that you have written a function to calculate the Euclidean distance between two points, recreate the function using default arguments of 0 for all inputs. Name your function `euclidean_distance_with_defaults`. The arguments should still be named `x1`, `y1`, `x2`, `y2`. Your function

Write your function in the cell below (that contains the ellipsis `...`).

Your function should be in the form:
``` PYTHON
def euclidean_distance_with_defaults(...):
    x = ...
    y = ...
    distance = ...
    return distance
```

```python
...
```

```python
euclidean_distance_with_defaults(1,1)
```

```python
grader.check("q3.1")
```

## Question 4

A ball is launched in the air with initial velocity $v_o$ from an initial position of coordinates $(x_0, y_0)$ as illustrated below: 

<center> <img src="resources/E7_Lab2_1.jpg" style="width: 300px;"/> </center>
<center> Figure 1: Ball launching at angle $\theta$ from the x-axis with inital velocity $v_0$ </center>

The ball is launched at an angle $\theta$ (expressed in degrees) from the x-axis. The
equations describing the motion (i.e. the position of the ball $(x(t), y(t)$) as a function
of time $t$) of the ball are:

$$ x(t) = x_0 + v_0 t \cos(\theta) $$
$$ y(t) = y_0 + v_0 t \sin(\theta) - \frac{1}{2}gt^2 $$

where $g$ is the gravitational acceleration ($g = 9.81 m / s^2$). Since $\theta$ is given in degrees,
you will want to convert it first into radians before using the `sin` or `cos` functions.

**Question 4.1:** Write a function that calculates the time at which the ball reaches the ground.

Note that when the ball hits the ground, $y = 0$. Also, before writing your code,
don't forget to use pen and paper to manipulate the equation(s) above and get a
mathematical expression for time.

```python
# run this cell to create a global variable g, but make sure you don't have any other global variables also named g
g = 9.81
```

```python
def proj_time(y0, v0, theta):
    b = ...
    time = ...
    return time
```

```python
proj_time(0, 15, 40)
```

```python
grader.check("q4.1")
```

**Question 4.2:** Write a function that calculates the horizontal distance traveled by the ball before
it reaches the ground. You can use the function `proj_time` and the equations of
motion.

```python
def proj_distance(y0, v0, theta):
    time = ...
    dist = ...
    return dist
```

```python
grader.check("q4.2")
```

```python
import matplotlib.image as mpimg
img = mpimg.imread('resources/animal.png')
imgplot = plt.imshow(img)
imgplot.axes.get_xaxis().set_visible(False)
imgplot.axes.get_yaxis().set_visible(False)
print("Congrats on finishing lab 2!")
plt.show()
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q1 (points: 2)
```
>>> assert get_hash(solve_quadratic(1, -7, 10)) == '5bb72c37ab645d2a79a1b1d0a32a17de'
>>> assert get_hash(solve_quadratic(5, -23, 2)) == '582fa9d143bfec5b7bfba577d5839f89'
```


Tests for q2.0 (points: 2)
```
>>> assert get_hash(c14_dating(10000)) == 'e5c2554f72f6c2498e4446e2da5a9508'
>>> assert get_hash(c14_dating(5730)) == '6d1a992c65604247a78a358d7ae42f0a'
```


Tests for q2.1 (points: 2)
```
>>> assert get_hash(exp_decay(10000,decay_constant=0.005)) == '66233a3eb22ff7c0c4e62df274cc5258'
>>> assert get_hash(exp_decay(5700,decay_constant=0.0005)) == '2bf78b853f1afdbfaf5b82123363ec06'
>>> assert get_hash(exp_decay(6000)) == '0cf1764baab2d028cd7a71aacd24b5ac' , 'check defult argument'
```


Tests for q3.0 (points: 2)
```
>>> assert get_hash(euclidean_distance(6,3,3,7)) == '336669dbe720233ed5577ddf81b653d3'
>>> assert get_hash(euclidean_distance(5.5,3.1,8.7,2.7)) == 'e704209f3a5cbd59f3c54004b9eb96e2'
```


Tests for q3.1 (points: 2)
```
>>> assert get_hash(euclidean_distance_with_defaults()) == get_hash(euclidean_distance_with_defaults(0,0,0,0))
>>> assert get_hash(euclidean_distance_with_defaults(y2=5,y1=4,x2=1,x1=3)) == '8c14832b68cbad6cd2cd54ee0a29d42f'
>>> assert get_hash(euclidean_distance_with_defaults(x2=4.5,y1=3.9)) == 'f34c28133868a9adf9b1dbe7277c5eb7'
```


Tests for q4.1 (points: 1)
```
>>> g = 9.81
>>> assert get_hash(proj_time(0, 15, 40)) == '085a4d0b759822dd9066e0f3eebe96b5'
```


Tests for q4.2 (points: 1)
```
>>> assert get_hash(proj_distance(0, 15, 40)) == '901c8912f3bd90f4c881badd02f571cb'
```
