```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

# E7: Introduction to Computer Programming for Scientists and Engineers

## Lab Assignment 2

For each question, you will have to fill in one or more Python functions. We provide an autograder with a number of test cases that you can use to test your function. Note that the fact that your function works for all test cases thus provided does necessarily guarantee
that it will work for all possible test cases relevant to the question. It is your responsibility
to test your function thoroughly, to ensure that it will also work in situations not covered
by the test cases provided

```python
# Please run this cell, and do not modify the contets
import math
import numpy as np
import matplotlib.pyplot as plt
#plt.style.use('fivethirtyeight')
#plt.rcParams["figure.figsize"] = (20,7)
%matplotlib inline

import hashlib
def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

## Question 1

Write a Python function that calculates the roots of a quadratic equation $ax^2+bx+
c = 0$ of unknown $x$, using the quadratic formula below. 

$$x = \frac{-b\pm \sqrt{b^2-4ac}}{2a}$$

The function output should be a two element tuple. The first element, `elem_1`, should be the root computed by setting the $\pm$ to $+$, and the second element, `elem_2`, should be the root computed by setting the $\pm$ to $-$.

The discriminant is $b^2 - 4ac$. You may assume $b^2 - 4ac \ge 0$.

We will cover tuples next week. For now, all you need to know is when you have multiple outputs from a function, Python returns all the values as a sequence of values called a tuple.

```python
# ANSWER CELL
def solve_quadratic(a, b, c):
    """
    Arguments
    ---------
    a : coefficient of x**2
    b : coefficient of x
    c : constant
    
    Returns
    -------
    elem_1 : larger root of quadratic equation
    elem_2 : smaller root of quadratic equation
    """
    discriminant = ...
    elem_1 = ...
    elem_2 = ...
    return elem_1, elem_2 
```

```python
solve_quadratic(6,-17,12)
```

```python
grader.check("q1")
```

## Question 2.0

Carbon dating is a method for determining the age of an object containing the radioisotope carbon-14. Since carbon-14 is subject to radioactive decay, the number of carbon-14 atoms in a sample will decay over time according to the following exponential decay law:

$$ N(t) = N_0 e^{-\lambda t} $$

where $N_0$ is the number of atoms of the isotope in the original sample, $N(t)$ is the remaining number of atoms after a time $t$, $\lambda$ is the expoential decay constant and $t$ is the time.

Write a function that calculates the remaining fraction of carbon-14 in a sample after a certain amount of time (in years). The remaining fraction is defined as the ratio of the number of atoms at the present time ($N(t)$) to the initial number of atoms ($N_0$). You can re-arrange the decay formula above to find an expression for this ratio.

Assume $\lambda = 0.00012097 \: \text{years}^{-1}$.

```python
# ANSWER CELL
def c14_dating(time):
    ratio = ...
    return ratio
```

Try out your new `c14_dating` function in the cell below!

You do not need to store any outputs from the cell below. It is to allow you to test your function.

```python

```

```python
grader.check("q2.0")
```

## Question 2.1

Now create a more general version of the function in Q2.0 and name the new function `exp_decay`. This new function will perform the same calculation as `c14_dating` but include a default argument for $\lambda$, the exponential decay constant. Call the argument for $\lambda$ `decay_constant`. Set the default value for `decay_constant` equal to the exponential decay constant for c-14.

Write your function in the cell below (that contains the ellipsis `...`).

Your function should be in the form:
``` PYTHON
def exp_decay(...):
    ratio = ...
    return ratio
```

```python
# ANSWER CELL
...
```

Try out your new `exp_decay` function in the cell below!

You do not need to store any outputs from the cell below. It is simply to allow you to test your function.

```python

```

```python
grader.check("q2.1")
```

## Question 3

A ball is launched in the air with initial velocity $v_o$ from an initial position of coordinates $(x_0, y_0)$ as illustrated below: 

<center> <img src="resources/E7_Lab2_1.jpg" style="width: 300px;"/> </center>
<center> Figure 1: Ball launching at angle $\theta$ from the x-axis with inital velocity $v_0$ </center>

The ball is launched at an angle $\theta$ (expressed in degrees) from the x-axis. The
equations describing the motion (i.e. the position of the ball $(x(t), y(t)$) as a function
of time $t$) of the ball are:

$$ x(t) = x_0 + v_0 t \cos(\theta) $$
$$ y(t) = y_0 + v_0 t \sin(\theta) - \frac{1}{2}gt^2 $$

where $g$ is the gravitational acceleration ($g = 9.81 m / s^2$). Since $\theta$ is given in degrees,
you will want to convert it first into radians before using the `sin` or `cos` functions.

**Question 3.0:** Write a function that calculates the time at which the ball reaches the ground.

Note that when the ball hits the ground, $y = 0$. Also, before writing your code,
don't forget to use pen and paper to manipulate the equation(s) above and get a
mathematical expression for time.

```python
# run this cell to create a global variable g, but make sure you don't have any other global variables also named g
g = 9.81
```

```python
# ANSWER CELL
def proj_time(y0, v0, theta):
    b = ...
    time = ...
    return time
```

Try out your new `proj_time` function in the cell below!

You do not need to store any outputs from the cell below. It is simply to allow you to test your function.

```python

```

```python
grader.check("q3.0")
```

**Question 3.1:** Write a function that calculates the horizontal distance traveled by the ball before
it reaches the ground. You can use the function `proj_time` and the equations of
motion.

```python
# ANSWER CELL
def proj_distance(y0, v0, theta):
    time = ...
    dist = ...
    return dist
```

Try out your new `proj_distance` function in the cell below!

You do not need to store any outputs from the cell below. It is simply to allow you to test your function.

```python

```

```python
grader.check("q3.1")
```

**Question 3.2:** Calculate the horizontal distance traveled by a ball starting at $5 m$ above the ground, with a velocity of 1 $m/s$ at an angle of 30 degrees above horizontal. Store the value as `q3_2`.

```python
# ANSWER CELL
q3_2 = ...
q3_2
```

```python
grader.check("q3.2")
```

**Question 3.3:** Reassign the gravitational acceleration variable, `g`, to 32.2 $ft/s^2$. Then calculate the horizontal distance traveled by a ball starting at $5 ft$ above the ground, with a velocity of 1 $ft/s$ at an angle of 30 degrees above horizontal. Store the value as `q3_3`.

```python
# ANSWER CELL
q3_3 = ...
q3_3
```

```python
# Do not remove. Run this cell to reassign g to 9.81 m/s^2.
g = 9.81
```

```python
grader.check("q3.3")
```

```python
import matplotlib.image as mpimg
img = mpimg.imread('resources/animal.png')
imgplot = plt.imshow(img)
imgplot.axes.get_xaxis().set_visible(False)
imgplot.axes.get_yaxis().set_visible(False)
print("Congrats on finishing lab 2!")
plt.show()
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: 5)
```
>>> assert get_hash(solve_quadratic(5, -23, 2)) == '582fa9d143bfec5b7bfba577d5839f89'
>>> assert get_hash(solve_quadratic(1,2,1)) == 'bb6dedb4a3fcd11f0b5570754ca0fc80'
>>> assert get_hash(np.round(solve_quadratic(5.2,3.5,-2.1),5)) == '6bf522b49df28612d661b9e4dca578a3'
```


Tests for q2.0 (points: 4)
```
>>> assert get_hash(round(c14_dating(5730),4)) == 'd310cb367d993fb6fb584b198a2fd72c'
>>> assert get_hash(round(c14_dating(0.956),4)) == '94eab0ef89d927c724320fec6169d935'
>>> assert get_hash(round(c14_dating(14023),4)) == 'adc61c91420dc18855e1b8f869d62f02'
```


Tests for q2.1 (points: 5)
```
>>> assert get_hash(exp_decay(10000,decay_constant=0.005)) == '66233a3eb22ff7c0c4e62df274cc5258'
>>> assert get_hash(exp_decay(5700,decay_constant=0.0005)) == '2bf78b853f1afdbfaf5b82123363ec06'
>>> assert get_hash(exp_decay(6000)) == '0cf1764baab2d028cd7a71aacd24b5ac' , 'check defult argument'
>>> assert get_hash(exp_decay(1,decay_constant=0.5)) == '7434e8620009ccbb8dbb5db6694a4c98'
>>> assert get_hash(exp_decay(106.34,decay_constant=0.0034)) == '6c27c52652a4a9432d93d4470f655cda'
>>> assert get_hash(exp_decay(1456)) == '9f3dd079930d3720fe56851394056889'
```


Tests for q3.0 (points: 4)
```
>>> g = 9.81
>>> assert get_hash(proj_time(0, 15, 40)) == '085a4d0b759822dd9066e0f3eebe96b5'
>>> assert get_hash(proj_time(10, 0, 45)) == '5d54d3cd4e1e02079de389a99a64931f'
```


Tests for q3.1 (points: 4)
```
>>> g = 9.81
>>> assert get_hash(proj_distance(0, 15, 40)) == '901c8912f3bd90f4c881badd02f571cb'
>>> assert get_hash(proj_distance(10, 0, 45)) == '30565a8911a6bb487e3745c0ea3c8224'
```


Tests for q3.2 (points: 1)
```
>>> assert get_hash(np.round(q3_2,4)) == 'c5204578f71b466f2818d519b95a925e'
```


Tests for q3.3 (points: 2)
```
>>> assert get_hash(np.round(q3_3,4)) == 'd65764cf7615df6f1cc880fdfb53c669'
```
