```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04.ipynb")
```

# E7: Introduction to Computer Programming for Scientists and Engineers

## Lab Assignment 4

For each question, you will have to fill in one or more Python functions. We provide an autograder with a number of test cases that you can use to test your function. Note that the fact that your function works for all test cases thus provided does necessarily guarantee that it will work for all possible test cases relevant to the question. It is your responsibility to test your function thoroughly, to ensure that it will also work in situations not covered by the test cases provided

## Score Breakdown
Question| Points
--- | --
1 | 5
2 | 6
3 | 7
4 | 7
Total | 25

```python
# Please run this cell, and do not modify the contets
import numpy as np
import math
import hashlib
import matplotlib.pyplot as plt
def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

## Question 1: Interest

The interest rate, $r$, on a principal, $P_0$, is a payment for allowing the bank to use an investor's money. Compound interest is accumulated according to the recursive formula $P_n = (1 + r)P_{n-1}$, where $P_n$ is the balance after n compounding periods (e.g. years), and $P_{n-1}$ is the balance after the previous period. The interest rate $r$ is given
in decimal, e.g. 0.1, which is 10%. 

Write a function `mySavingsPlan(P0, rate, goal)` with output `years`. `years` is the number of years it will  require for an investor to reach the investment goal defined by `goal` while investing an original amount `P0` at the annual interest rate, and specified by the variable `rate`. For the purpose of this problem, consider the investment balance to change only once per year (i.e. the output years should be a whole number).

*Hint:*

```python
# ANSWER CELL
def mySavingsPlan(p0, rate, goal):
    ...
    return years
```

Try out your new `mySavingsPlan` function in the cell below for `p0 = 1000` dollars, `rate = 0.05`, and `goal = 10000` dollars! Feel free to try it for other input times.

You do not need to store any outputs from the cell below. It is to allow you to test your function.

```python

```

```python
grader.check("q1")
```

## Question 2: Searching

Write a Python function `myMin(P)` with output `value, location`, where `P` is a class double (numeric) array. The output value is the minimum value in the array `P` and location is the index of the minimum value. Note that the minimum
value could appear more than once. If this is the case, location should specify the
indices of all occurrences of the minimum value. You may NOT use the built-in
Python command `min` in your function.

```python
# ANSWER CELL
def myMin(P):
    """
    Examples
    --------
    (1, 6) = myMin([5,4,5,2,6,7,1,8,9])
    (2, [0, 4, 7]) = myMin([2,3,5,7,2,4,3,2])
    (-7, 3) = myMin([2,-3,5,-7,2,4,3,2])
    (-3, [1, 6]) = myMin([2,-3,5,7,2,4,-3,2])
    
    Parameters
    ----------
    P: a float array
    
    Returns
    -------
    value: min value in P
    location: index of min value(s)
    """
    ...
    return value, location
    
```

Try out your new `myMin` function in the cell below for ` P = [2,-3,5,7,2,4,-3,2]`. Feel free to try it for other input times.

You do not need to store any outputs from the cell below. It is to allow you to test your function.

```python

```

```python
grader.check("q2")
```

## Question 3: Rainfall Totals

Design a function that computes the total rainfall at a location using daily rainfall
measurements entered by a user. The list may contain the number -999 indicating
the end of the data from a given period. There may be negative numbers other than
-999 in the list which represent erroneous measurements which should be neglected.
Produce the total and average of the non-negative values in the list up to the first -999
(if there is a -999 value). 

Your function `myRainfall(rainData)` should return `rainTot, rainAvg`,
where `rainData` is a double array of daily rainfall data, `rainTot` is the total rainfall
in the period represented by either the entire array `rainData` or the period from the
beginning until the first occurrence of -999. `rainAvg` is the average daily rainfall during
this period (not counting days with neglected measurements).

```python
# ANSWER CELL
def myRainfall(rainData):
    """
    Examples 
    --------
    (22, 4.4) = myRainfall([2,4,-5,1,7,-2,8,-999,4,6,4])
    (36, 4.0) = myRainfall([2,4,-5,1,7,-2,8,0,4,6,4])
    (6, 3.0) = myRainfall([2,4,-999,1,7,-2,8,-999,4,6,4])
    (33.4, 4.771428571428571) = myRainfall([2.3,4.1,-5.4,1.0,7.1,-2,8.0,4.9,6,-999])
    (0, 0.0) = myRainfall([0,-999,-5,1,7,-2,8,-999,4,6,2])
    
    Parameters
    ----------
    rainData: a float array, with positive values, negative values, or -999
    
    Returns
    -------
    total rainfall: float, summation of non-negative values until the first occurrence of -999
    average rainfall : float, total calculated rainfall over recorded days
    """
    ...
    return rainTot, rainTot / float(count_days)
```

Try out your new `myRainfall` function in the cell below for ` rainData = [2,4,-5,1,7,-2,8,-999,4,6,4]`. Feel free to try it for other input values.

You do not need to store any outputs from the cell below. It is to allow you to test your function.

```python

```

```python
grader.check("q3")
```

## Question 4: Summation Target

Write a function that takes an array of $K$ integers and a single desired value to serve as a target summation. Are there three elements $a, b, c$ in the array such that $a + b + c$ = `sumtarget`? Find all unique triplets in the array which give the sum of the target summation. Note: Each row in the solution should be listed in ascending order and no duplicate rows should be listed.

Your function `mySummation(array, sumtarget)` should return `results`, where `results` is an Mx3 double array with all the 1x3 arrays of values that add up to `sumtarget`, `array` is an array of numbers in ascending order, and `sumtarget` is the desired sum of any three elements of `array`.

*Hint:* You may find `numpy.vstack` and `numpy.empty` useful for creating your `results` array, check out `numpy.vstack` documentation [here](https://numpy.org/doc/stable/reference/generated/numpy.vstack.html) `numpy.empty` [here](https://numpy.org/doc/stable/reference/generated/numpy.empty.html).

```python
# ANSWER CELL
def mySummation(array, sumtarget):
    """
    Examples
    --------
    array = [1, 1, 5, 6, 7, 8, 10]
    [[1 1 8]] = mySummation(array, 10)
        
    array = [0, 1, 1, 2, 3, 4]
    [[0 1 4],[0 2 3],[1 1 3]] = mySummation(array, 5)
        
    array = [0, 1, 1, 2, 3, 4]
    [] = mySummation(array, 10)
    
    Parameters
    ----------
    array: array of numbers in ascending order
    sumtarget: desired sum of any three elements of input array
    
    Returns
    -------
    results: Mx3 double array with all the unique 1x3 arrays of values that add up to sumtarget 
    """
    ...
    return results
```

Try out your new `mySavingsPlan` function in the cell below for `array = [0, 1, 1, 2, 3, 4]` and `sumtarget = 10`! Feel free to try it for other input values.

You do not need to store any outputs from the cell below. It is to allow you to test your function.

```python

```

```python
grader.check("q4")
```

```python
import matplotlib.image as mpimg
img = mpimg.imread('resources/animal.png')
imgplot = plt.imshow(img)
imgplot.axes.get_xaxis().set_visible(False)
imgplot.axes.get_yaxis().set_visible(False)
print("Congrats on finishing lab 4 on branching!")
plt.show()
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q1 (points: 5)
```
>>> assert get_hash(mySavingsPlan(2000, 0.15, 10000)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(mySavingsPlan(4000, 0.10, 300)) == 'cfcd208495d565ef66e7dff9f98764da', "What if p0 is greater than goal? Are additional years needed for saving?"
```
```
>>> assert get_hash(mySavingsPlan(1000, 0.05, 10000)) == '642e92efb79421734881b53e1e1b18b6'
```


Tests for q2 (points: 6)
```
>>> assert get_hash(myMin([2])[-1]) == 'cfcd208495d565ef66e7dff9f98764da'
>>> assert get_hash(myMin([38,8,9,77,12,2])) == '6693afa515ec41da1d119faeb81df840'
```


Tests for q3 (points: 7)
```
>>> assert get_hash(myRainfall([2.3,4.1,-5.4,1.0,7.1,-2,8.0,4.9,6,-999])) == 'e0e5d666f240241f3ddcf3af9caf411c'
```
```
>>> assert get_hash(myRainfall([2,1.3,-99,33,0.99,-999])[0]) == '652126365f55aa7ccf24e90ed647c3a3'
>>> assert get_hash(myRainfall([2,1.3,-99,33,0.99,-999])[1]) == 'a311e6b63f7e2eb2e08f75cb9d5c882c'
```


Tests for q4 (points: 7)
```
>>> assert get_hash(mySummation([0, 1, 1, 2, 3, 4],10)) == 'd751713988987e9331980363e24189ce'
>>> assert get_hash(int(sum(mySummation([0, 0, 2, 4, 6, 8, 14],14)[:,2]))) == '34173cb38f07f89ddbebc2ac9128303f'
```
