```python
# Initialize Otter
import otter
grader = otter.Notebook("lab03.ipynb")
```

# E7: Introduction to Computer Programming for Scientists and Engineers

## Lab Assignment 3

For each question, you will have to fill in one or more Python functions. We provide an autograder with a number of test cases that you can use to test your function. Note that the fact that your function works for all test cases thus provided does necessarily guarantee that it will work for all possible test cases relevant to the question. It is your responsibility to test your function thoroughly, to ensure that it will also work in situations not covered by the test cases provided.

## Score Breakdown
Question| Points
--- | --
  1 | 1
2.0 - 2.1 | 2
3.0 - 3.1 | 4
4.0 - 4.3 | 4
5 | 3
6.0 | 5
6.1 | 1
Total | 20

```python
# Please run this cell, and do not modify the contents
import numpy as np
import math
import matplotlib.pyplot as plt
import re
np.seterr(all='ignore');
import hashlib
def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

## Introduction
For this lab assignment, you will analyze water storage of three reservoirs in California from Jan-Dec 2022. You will clean and manipulate the data in questions 1-4. Then you will calculate basic statistics for each reservoir in questions 5 and 6.

### Background
The three reservoirs you will analyze are Lake Oroville, Shasta Lake, and Hetch Hetchy Reservoir. Shasta Lake and Lake Oroville are the largest and second-largest reservoirs in California, providing flood-control, hydropower, and water storage for the state. Shasta Lake is a key reservoir within the federal government's Central Valley Project. Similarly, Lake Oroville is a major reservoir in the state government's California State Water Project. Hetch Hetchy reservoir is within Yosemite National Park and provides water to the City of San Francisco and other counties in the San Francisco Bay Area.

### Data
We will read a text file that includes the monthly storage volumes for each reservoir, taken from January through December 2022, in units of acre-feet. We will provide you with the command to read the file, but in the future, you will learn how to read text files and other file types. In the code cell below, we are reading a text file and saving its values as `reservoir_data`. Run the cell below and do not modify its contents.

```python
# Please run this cell, and do not modify the contents
reservoir_data = np.loadtxt('resources/res_vals_int.csv',delimiter='\t')
print(reservoir_data)
```

The 12 rows correspond to the monthly storage volumes and the three columns correspond to the three reservoirs in the following order:
1. Lake Oroville
2. Shasta Lake
3. Hetch Hetchy reservoir

In the `reservoir_names` list, you have the names of each reservoir as they appear in the `reservoir_data` array, but with some missing/incomplete entires. In this assignment, you will manipulate the `reservoir_names` list to match your data. For now, run the code cell below without modifying its contents.

```python
# Please run this cell, and do not modify the contents
reservoir_names = ['Lake Oroville','Shasta']
print(reservoir_names)
```

## Question 1
In the list `reservoir_names`, Shasta Lake is listed as just `Shasta`. Change the value from `Shasta` to `Shasta Lake` in the list. Do not create a new variable.

```python
# ANSWER CELL
...
reservoir_names
```

```python
grader.check("q1")
```

## Question 2.0
Our list `reservoir_names` only has two entries! We are missing the third entry `Hetch Hetchy Reservoir`. Add `Hetch Hetchy Reservoir` to the end of the `reservoir_names` list. Do not redefine the list. Use appropriate list methods that allow you to add an entry at the end.

```python
# ANSWER CELL
...
reservoir_names
```

```python
grader.check("q2.0")
```

## Question 2.1
Now that you have the correct values in `reservoir_names`, check to see if certain reservoirs are in your list. Check if "Pardee Reservoir" is in the list of reservoirs using list methods. Write your boolean statement in `q2_1`.

*Note:* Pardee Reservoir is a major reservoir for East Bay Municipal Utility District.

```python
# ANSWER CELL
q2_1 = ...
```

```python
grader.check("q2.1")
```

## Question 3
In this question, you will create a range of numbers 1-12 to represent the rows within the `reservoir_data` array (i.e., the months Jan - Dec).

## Question 3.0
Create range of integers from 1 to 12, inclusive, using the `range` function. Store the range of numbers as `months_range`.

```python
# ANSWER CELL
...
```

```python
grader.check("q3.0")
```

## Question 3.1
Convert `months_range` to a numpy array and to a list. Store the values as `months_nparray` and `months_list` respectively. The variables `months_nparray` and `months_list` should contain the integers 1 through 12, just like `months_range`.

```python
# ANSWER CELL
months_nparray = ...
months_list = ...
print(months_nparray , months_list)
```

```python
grader.check("q3.1")
```

## Question 4
Now you will start to analyze the reservoir data. For this question, do not edit `reservoir_data`. Refer to the instructions at the beginning of the assignment on which column of the data corresponds to which reservoir.

## Question 4.0
First, slice the array to select the column of data of Lake Oroville.

Store the sliced data as `OrovilleData`.

```python
# ANSWER CELL
...
```

```python
grader.check("q4.0")
```

## Question 4.1
Slice the array to select the winter storage volumes for Hetch Hetchy Reservoir. For this question, use January, February, and March as the winter months.

Store the sliced data as `WinterHetchHetchy`.

```python
# ANSWER CELL
...
```

```python
grader.check("q4.1")
```

## Question 4.2
Find the mean storage volume for Summer (July through September) for each reservoir. California receives very little rainfall in the summer months (July through September), so reservoir storage during these months is of particular interest.

Store the mean values as an array called `SummerMeanStorage`.

*Hint*: Slice the data to the months of interest, then calculate the mean across the appropriate axis.

```python
# ANSWER CELL
...
```

```python
grader.check("q4.2")
```

## Question 4.3
How much more water was stored in Lake Oroville than in Hetch Hetchy in the month of April? Report your answer in thousand acre-feet (TAF).

Store the difference in thousand acre-feet as `DiffApril`.

*Hint*: Divide by 1000 to convert from acre-feet to TAF.

```python
# ANSWER CELL
...
```

```python
grader.check("q4.3")
```

## Question 5
Write a function called `ReservoirInfo` which takes information about reservoir storage as the inputs and outputs a statement: ‘In [Month], [Name] held [Volume] acre-feet of water.’ where the words in brackets are the input month, reservoir name, and storage volume, respectively. An example output for would be: ‘In January, Lake Oroville held 1642000 acre-feet of water.’

Your input variables should be called `name` , `month`, and `volume`. You may have your inputs in any order.

Your function should return, but not print, your statement.

```python
# ANSWER CELL
...
```

```python
grader.check("q5")
```

## Question 6.0
Finally for this assignment, you will define the input data and title string for a pre-made plotting function called `plot_reservoir_storage`. The `plot_reservoir_storage` function has two inputs:
- `res_name` is the reservoir's name as a string
- `res_data` is the storage data of the reservoir (for only one reservoir)

For this question, define `title_string` as a format string to have the `res_name` input followed by "storage in 2022". For example, a graph of Lake Oroville should have the title "Lake Oroville storage in 2022". The variable `title_string` should automatically update based on the input `res_name`. Be careful with your string formatting. Any small modification, such as extra space or character, will result in a failed test.

Additionally, define `x_label` and `y_label` as format strings for the axis labels of your plot. Your labels should include the correct units "month" and "acre-feet". Be careful with your string formatting. Any misspelling will result in a failed test.

You will plot Hetch Hetchy data using the `plot_reservoir_storage` function.
Define `name_input` as the value from the `reservoir_names` list corresponding to Hetch Hetchy. 
Define `data_input` as an array slice from `reservoir_data` corresponding to Hetch Hetchy storage data.

Finally, call the function `plot_reservoir_storage` and pass the inputs `name_input` and `data_input`. Assign the function call to `q6_0`.

*Note:* This question relies on the completion of Q2.0 and Q3.0.

```python
# ANSWER CELL
def plot_reservoir_storage(res_name,res_data):
    title_string = ...
    x_label = ...
    y_label = ...
    
    # Do not edit the function code below this line.
    plt.plot(months_range,res_data)
    plt.xlabel(x_label)
    plt.ylabel(y_label)
    plt.title(title_string)
    plt.show()
    return title_string , x_label , y_label

name_input = ...
data_input = ...

# Run the below code, do not edit
q6_0 = ...
```

```python
grader.check("q6.0")
```

## Question 6.1
In Q6.0, all three outputs from the `plot_reservoir_storage` function were saved as a single variable called `q6_0`. 

Evaluate the type of `q6_0` and save the type as `q6_1a`.\
Evaluate the length of `q6_0` and save the length as `q6_1b`.

```python
# ANSWER CELL
q6_1a = ...
q6_1b = ...
print(f'Type: {q6_1a} \nLength: {q6_1b}')
```

```python
grader.check("q6.1")
```

```python
import matplotlib.image as mpimg
img = mpimg.imread('resources/animal.jpg')
imgplot = plt.imshow(img)
imgplot.axes.get_xaxis().set_visible(False)
imgplot.axes.get_yaxis().set_visible(False)
print("Congrats on finishing lab 3!")
plt.show()
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: 1)
```
>>> assert get_hash(len(reservoir_names)) == 'c81e728d9d4c2f636f067f89cc14862c' , 'You should only have two values (for now).'
>>> assert get_hash(reservoir_names[0]) == 'b947db2fb4b5855984587a0d4fabf896' , 'Do not edit the Lake Oroville entry'
>>> assert get_hash(reservoir_names[1].upper()) =='5f3b912d347dbbe2de8a9db5c1226e44' , 'Check value of Shasta Lake entry'
```


Tests for q2.0 (points: 1)
```
>>> assert get_hash(len(reservoir_names)) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3' , 'Check length of list'
>>> assert get_hash(reservoir_names[-1].upper()) == 'e5bb3ed8fc109fa9f05b0d538a30edf8' , 'Check value Hetch Hetchy entry'
```


Tests for q2.1 (points: 1)
```
>>> assert get_hash(q2_1) == 'f8320b26d30ab433c5a54546d21f414c'
```


Tests for q3.0 (points: 2)
```
>>> assert get_hash(months_range) == '5319bc85f3a29d6cd7b6b65df958ae78'
>>> assert get_hash(months_range[0]) == 'c4ca4238a0b923820dcc509a6f75849b' , 'Check start value'
>>> assert get_hash(months_range[-1]) == 'c20ad4d76fe97759aa27a0c99bff6710' , 'Check end value'
>>> assert get_hash(sum(months_range)) == '35f4a8d465e6e1edc05f3d8ab658c551' , 'Check step size'
```


Tests for q3.1 (points: 2)
```
>>> assert get_hash(type(months_nparray)) == '30b39aca73d0815f5ca1a06af1412947'
>>> assert get_hash(len(months_nparray)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(sum(months_nparray)) == '35f4a8d465e6e1edc05f3d8ab658c551'
```
```
>>> assert get_hash(type(months_list)) == '425ed5b515f0a67c6316f0ec499f86bf'
>>> assert get_hash(len(months_list)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(sum(months_list)) == '35f4a8d465e6e1edc05f3d8ab658c551'
```


Tests for q4.0 (points: 1)
```
>>> assert get_hash(len(OrovilleData)) == 'c20ad4d76fe97759aa27a0c99bff6710' , 'Check length of selection'
>>> assert get_hash(int(sum(OrovilleData))) == '97c2a51b4ddc77196fa563cd72957c1e' , 'Check which column corresponds to Lake Oroville'
>>> assert get_hash(OrovilleData.astype(int)) == 'd57f3682934eebd41c73b68817c6592d'
```


Tests for q4.1 (points: 1)
```
>>> assert get_hash(len(WinterHetchHetchy)) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3' , 'Check length of selection'
>>> assert get_hash(int(sum(WinterHetchHetchy))) == '50037bd6578a3e4db86ac430e190f7cb' , 'Check which column corresponds to Lake Oroville'
>>> assert get_hash(WinterHetchHetchy.astype(int)) == '7b67ebf4e3d0d34b2c62dba86b620e1b'
```


Tests for q4.2 (points: 1)
```
>>> assert get_hash(len(SummerMeanStorage)) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3' , 'Check length of selection'
>>> assert get_hash(round(sum(SummerMeanStorage),-2)) == 'd704e6175c0f8f1ea3d701b0b6d07b88' , 'Check slice selection'
>>> assert get_hash(round(SummerMeanStorage[1],-2)) == '71f214fd332f17d09f2f72f4dbbd29ec' , 'Check mean calculation'
```


Tests for q4.3 (points: 1)
```
>>> assert get_hash(round(DiffApril)) == '8f19793b2671094e63a15ab883d50137'
```


Tests for q5 (points: 3)
```
>>> assert get_hash(len(ReservoirInfo(name=reservoir_names[2],month='February',volume=reservoir_data[1,2]))) == '735b90b4568125ed6c3f678819b6e058' , 'Check spacing and data type'
```
```
>>> assert get_hash('FAKEMONTH' in ReservoirInfo(name='PLACEHOLDER',month='FAKEMONTH',volume=0)) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash('PLACEHOLDER' in ReservoirInfo(name='PLACEHOLDER',month='FAKEMONTH',volume=0)) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash('0' in ReservoirInfo(name='PLACEHOLDER',month='FAKEMONTH',volume=0)) == 'f827cf462f62848df37c5e1e94a4da74'
```


Tests for q6.0 (points: 5)
```
>>> assert get_hash(name_input) == 'f87798ca0a0f836f16995cc5c7955c2a'
>>> assert get_hash(q6_0[0]) == '75cbfe9a1d825c40fd979cc5614764aa' , 'Check title string'
>>> assert get_hash("MONTH" in q6_0[1].upper()) == 'f827cf462f62848df37c5e1e94a4da74' , 'Check x-label'
>>> # Uses regex search to look for acre feet. Any character allowed between acre and feet. Case insensitive.
>>> assert re.search(r'ACRE.FEET',q6_0[2].upper()) , 'Check y-label'
>>> assert get_hash(data_input.astype(int)) == 'ac3c5f2b1ab399ae627808b2459456ca'
```


Tests for q6.1 (points: 1)
```
>>> assert get_hash(q6_1a) == '22b3afce62075c7012f8e5041adfee16'
>>> assert get_hash(q6_1b) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3'
```
