```python
# Initialize Otter
import otter
grader = otter.Notebook("lab03_redo.ipynb")
```

# E7: Introduction to Computer Programming for Scientists and Engineers

## Lab Assignment 3

For each question, you will have to fill in one or more Python functions. We provide an autograder with a number of test cases that you can use to test your function. Note that the fact that your function works for all test cases thus provided does necessarily guarantee that it will work for all possible test cases relevant to the question. It is your responsibility to test your function thoroughly, to ensure that it will also work in situations not covered by the test cases provided.

## Score Breakdown
Question| Points
--- | --
  1 | 1
2.0 | 1
2.1 | 1
3.0 | 1
3.1 | 2
3.2 | 2
  4 | 3
5.0 | 1
5.1 | 8
  6 | 5
Total | 10 | 25

```python
# Please run this cell, and do not modify the contents
import numpy as np
import math
import matplotlib.pyplot as plt
np.seterr(all='ignore');
import hashlib
def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

## Introduction
For this lab assignment, you will analyze water storage of three reservoirs in California from Jan-Dec 2022. You will clean and manipulate the data in questions 1-4. Then you will calculate basic statistics for each reservoir in questions 5 and 6.

### Background
The three reservoirs you will analyze are Lake Oroville, Shasta Lake, and Hetch Hetchy Reservoir. Shasta Lake and Lake Oroville are the largest and second-largest reservoirs in California, providing flood-control and water storage for the state. Shasta Lake is a key reservoir within the federal government's Central Valley Project. Similarly, Lake Oroville is a major reservoir in the the state government's California State Water Project. Hetch Hetchy reservoir is within Yosemite National Park and provides water to the City of San Francisco and other counties in the San Francisco Bay Area.

### Data
In the `reservoir_data_original` array, you have 12 reservoir storage values for each reservoir, taken monthly from January through December 2022, in units of acre-feet. The three columns of the `reservoir_data_original` array are Lake Oroville, Shasta Lake, and Hetch Hetchy reservoir, in this order. In this assignment, you will manipulate the data before analyzing each time series.

In the `reservoir_names` list, you have the names of each reservoir as they appear in the `reservoir_data_original` array. In this assignmen, you will manipulate the `reservoir_names` list to match your data.

```python
# Please run this cell, and do not modify the contents
reservoir_names = ['Lake Oroville','Shasta']
reservoir_data_original = np.loadtxt('resources/res_vals_str.csv',delimiter='\t',dtype=str)
```

## Question 1
`reservoir_names` lists Shasta Lake as just `Shasta`.In the list `reservoir_names`, change the value of the `Shasta` item to `Shasta Lake`. Do not create a new variable.

```python

```

```python
grader.check("q1")
```

## Question 2.0
Our list `reservoir_names` only has two entries! We are missing the third entry `Hetch Hetchy Reservoir`. Add `Hetch Hetchy Reservoir` to the end of the `reservoir_names` list.

```python
reservoir_names
```

```python
grader.check("q2.0")
```

## Question 2.1
Now that you have the correct values in `reservoir_names`, check to see if certain reservoirs are in your list. Check if "Pardee Reservoir" is in the list of reservoirs using list methods. Write your boolean statement in `q2_1`.

*Note:* Pardee Reservoir is a major reservoir for East Bay Municipal Utility District.

```python
q2_1 = ...
```

```python
grader.check("q2.1")
```

## Question 3
In this question, you will create a range of numbers 1-12 to represent rows the rows within the `reservoir_data` array (aka the months Jan - Dec).

## Question 3.0
Create range of integers from 1 to 12, inclusive, using the `range` function. Store the range of numbers as `months_range`.

```python
months_range = ...
months_range
```

```python
grader.check("q3.0")
```

## Question 3.1
Convert `months_range` to a numpy array and to a list. Store the values as `months_nparray` and `months_list` respectively. The variables `months_nparray` and `months_list` should contain the integers 1 through 12, just like `months_range`.

```python
months_nparray = ...
months_list = ...
print(months_nparray , months_list)
```

```python
grader.check("q3.1")
```

## Question 3.2
Now you have three versions of the numbers 1-12: `months_range`, `months_nparray`, and `months_list`. Use the function `getsizeof` to find the size of each of the three variables. Store the size in bytes of `months_range`, `months_nparray`, and `months_list` in `size_of_range`, `size_of_nparray`, and `size_of_list` respectively.

*Hint:* Check out the documentation for `sys.getsizeof` [here](https://docs.python.org/3/library/sys.html#sys.getsizeof).

```python
from sys import getsizeof
```

```python
size_of_range = ...
size_of_nparray = ...
size_of_list = ...
print(size_of_range , size_of_nparray , size_of_list)
```

```python
grader.check("q3.2")
```

## Question 4
Now look at the `reservoir_data_original` array. Unfortunately, the entries are all strings!

```python
print(reservoir_data_original)
```

Create a function called `str2int_function` to convert a single string entry to an integer. The input to this function should be a string, and the output should be an integer.

Once you create the `str2int_function` function, pre-made code will apply your function to each value in the `reservoir_data_original` array and save the new array as `reservoir_data`.

*Note:* There is a blank `str2int_function` function with placeholder inputs and outputs, so your code does not break if you run the notebook while empty.

*Hint:* `int()` can convert a string of numbers to integer, but cannot handle commas in the string.

```python
def str2int_function(placeholder):
    ...
    return placeholder

# Run the below code, do not edit
reservoir_data = np.vectorize(str2int_function)(reservoir_data_original)
reservoir_data
```

```python
grader.check("q4")
```

## Question 5.0
Now that you have changed `reservoir_data` to numbers, you will start to analyze the reservoir data. First, slice the array to select the column of data of Lake Oroville. Do not edit `reservoir_data`.

Store the sliced data as `Oroville_data`.

```python
Oroville_data = ...
Oroville_data
```

```python
grader.check("q5.0")
```

## Question 5.1
Fill in the `reservoir_change` function to print out the mean reservoir storage value, as well as January's and December's storage values.

The input to the `reservoir_change` function is a column of data from one reservoir location (such as `Oroville_data`). The outputs are the mean storage value, a string describing January's storage value, and a string describing December's storage value.

You will create the formatted strings to display January's storage value in `jan_storage_string` and to display December's storage value in `dec_storage_string`. The strings you create should meet the following criteria:
- Contain the correct storage value
- Have the units "AF" (for acre-feet) after the value
- Have the phrase "January reservoir storage" or "December reservoir storage" before the value.

*Hint:* January is the first value in the input array, and December is the last value.

```python
def reservoir_change(location_data):
    mean_storage = ...
    print('Mean reservoir storage:',int(mean_storage),' AF')
    jan_storage_string = ...
    print(jan_storage_string)
    dec_storage_string = ...
    print(dec_storage_string)
    return mean_storage , jan_storage_string , dec_storage_string

# Run the below code, do not edit
oroville_stats = reservoir_change(Oroville_data)
```

```python
grader.check("q5.1")
```

## Question 5.2
In Q5.1, all three outputs from the `reservoir_change` function were saved as a single variable called `oroville_stats`. 

Evaluate the type of `oroville_change` and save the type as `q5_2a`.\
Evaluate the length of `oroville_change` and save the length as `q5_2b`.

```python
q5_2a = ...
q5_2b = ...
print(q5_2a , q5_2b)
```

```python
grader.check("q5.2")
```

## Question 6
Finally for this assignment, you will define the input data and title string for a pre-made plotting function called `plot_reservoir_storage`. The `plot_reservoir_storage` function has two inputs:
- `res_name` is the reservoir's name as a string
- `res_data` is the storage data of a single reservoir.

For this question, define `title_string` to have the `res_name` input followed by "storage in 2022". For example, a graph of Lake Oroville should have the title "Lake Oroville storage in 2022".

You will plot Hetch Hetchy data using the `plot_reservoir_storage` function and will define the inputs into the function. Define `name_input` as the value from the `reservoir_names` list for Hetch Hetchy. And define `data_input` as an array slice from `reservoir_data` corresponding to Hetch Hetchy.

*Note:* This question relies on the completion of Q2.0, Q3.0, and Q4.

```python
def plot_reservoir_storage(res_name,res_data):
    title_string = ...
    
    # Do not edit the function code below this line.
    plt.plot(months_range,res_data)
    plt.xlabel('Month')
    plt.ylabel('Reservoir storage (acre-feet)')
    plt.title(title_string)
    plt.show()
    return title_string
name_input = ...
data_input = ...

# Run the below code, do not edit
q6 = plot_reservoir_storage(name_input,data_input)
```

```python
grader.check("q6")
```

```python
import matplotlib.image as mpimg
img = mpimg.imread('resources/animal.jpg')
imgplot = plt.imshow(img)
imgplot.axes.get_xaxis().set_visible(False)
imgplot.axes.get_yaxis().set_visible(False)
print("Congrats on finishing lab 3!")
plt.show()
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: 1)
```
>>> assert get_hash(len(reservoir_names)) == 'c81e728d9d4c2f636f067f89cc14862c' , 'You should only have two values (for now).'
>>> assert get_hash(reservoir_names[0]) == 'b947db2fb4b5855984587a0d4fabf896' , 'Do not edit the Lake Oroville entry'
>>> assert get_hash(reservoir_names[1].upper()) =='5f3b912d347dbbe2de8a9db5c1226e44' , 'Check value of Shasta Lake entry'
```


Tests for q2.0 (points: 1)
```
>>> assert get_hash(len(reservoir_names)) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3' , 'Check length of list'
>>> assert get_hash(reservoir_names[-1].upper()) == 'e5bb3ed8fc109fa9f05b0d538a30edf8' , 'Check value Hetch Hetchy entry'
```


Tests for q2.1 (points: 1)
```
>>> assert get_hash(q2_1) == 'f8320b26d30ab433c5a54546d21f414c'
```


Tests for q3.0 (points: 1)
```
>>> assert get_hash(months_range) == '5319bc85f3a29d6cd7b6b65df958ae78'
>>> assert get_hash(months_range[0]) == 'c4ca4238a0b923820dcc509a6f75849b' , 'Check start value'
>>> assert get_hash(months_range[-1]) == 'c20ad4d76fe97759aa27a0c99bff6710' , 'Check end value'
>>> assert get_hash(sum(months_range)) == '35f4a8d465e6e1edc05f3d8ab658c551' , 'Check step size'
```


Tests for q3.1 (points: 2)
```
>>> assert get_hash(type(months_nparray)) == '30b39aca73d0815f5ca1a06af1412947'
>>> assert get_hash(len(months_nparray)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(sum(months_nparray)) == '35f4a8d465e6e1edc05f3d8ab658c551'
```
```
>>> assert get_hash(type(months_list)) == '425ed5b515f0a67c6316f0ec499f86bf'
>>> assert get_hash(len(months_list)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(sum(months_list)) == '35f4a8d465e6e1edc05f3d8ab658c551'
```


Tests for q3.2 (points: 2)
```
>>> assert get_hash(size_of_range) == '642e92efb79421734881b53e1e1b18b6'
>>> assert get_hash(size_of_nparray) == '3644a684f98ea8fe223c713b77189a77'
>>> assert get_hash(size_of_list) == '37a749d808e46495a8da1e5352d03cae'
```


Tests for q4 (points: 3)
```
>>> assert get_hash(reservoir_data.size) == '19ca14e7ea6328a42e0eb13d585e4c22' , 'Check shape of array'
>>> assert get_hash(sum(sum(reservoir_data))) == 'f7a4bd529c2762d35ea70706a2b804eb' , 'Check values of array'
>>> assert get_hash(reservoir_data) == 'a312845f72761ce35271d3b4b168c6c1'
>>> assert get_hash(reservoir_data.dtype) == 'ff9b3f96d37353c528517bc3656a00a8' , 'Check dtype of array'
```


Tests for q5.0 (points: 1)
```
>>> assert get_hash(len(Oroville_data)) == 'c20ad4d76fe97759aa27a0c99bff6710' , 'Check length of selection'
>>> assert get_hash(sum(Oroville_data)) == '97c2a51b4ddc77196fa563cd72957c1e' , 'Check which column corresponds to Lake Oroville'
>>> assert get_hash(Oroville_data) == 'd57f3682934eebd41c73b68817c6592d'
```


Tests for q5.1 (points: 8)
```
>>> assert get_hash(int(oroville_stats[0])) == '940e9e08f28abe5dbb1a1518cd1210c9' , 'Check mean calculation or input values'
```
```
>>> assert get_hash('January reservoir storage' in oroville_stats[1]) == 'f827cf462f62848df37c5e1e94a4da74' , 'Check January string'
>>> assert get_hash('1642000' in oroville_stats[1]) == 'f827cf462f62848df37c5e1e94a4da74' , 'Check January value'
>>> assert get_hash('AF' in oroville_stats[1]) == 'f827cf462f62848df37c5e1e94a4da74' , 'Check January string'
```
```
>>> assert get_hash('December reservoir storage' in oroville_stats[2]) == 'f827cf462f62848df37c5e1e94a4da74' , 'Check December string'
>>> assert get_hash('1273042' in oroville_stats[2]) == 'f827cf462f62848df37c5e1e94a4da74' , 'Check December value'
>>> assert get_hash('AF' in oroville_stats[2]) == 'f827cf462f62848df37c5e1e94a4da74' , 'Check December string'
```
```
>>> assert get_hash(int(reservoir_change(reservoir_data[:,1])[0])) == '65035f8862f37a3103cfee9f3930692d' , 'Check reproducibility'
Mean reservoir storage: 1630893  AF
January reservoir storage: 1621000 AF
December reservoir storage: 1523676 AF
```


Tests for q5.2 (points: 1)
```
>>> assert get_hash(q5_2a) == '22b3afce62075c7012f8e5041adfee16'
>>> assert get_hash(q5_2b) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3'
```


Tests for q6 (points: 5)
```
>>> assert get_hash(name_input) == 'f87798ca0a0f836f16995cc5c7955c2a'
>>> assert get_hash(q6) == '75cbfe9a1d825c40fd979cc5614764aa'
>>> assert get_hash(data_input) == 'ac3c5f2b1ab399ae627808b2459456ca'
```
