```python
# Initialize Otter
import otter
grader = otter.Notebook("lab06.ipynb")
```

# E7: Lab Assignment 06

You must submit the lab to Gradescope by the due date. You will submit the zip file produced by running the final cell of the assignment.

## About this Lab
The objective of this assignment is to familiarize you with plotting as well as reading and writing data.

## Autograder
You will be provided with some test cases (i.e., sample input data) and associated answers (i.e., expected outputs) that you can use to help check your code. The provided test cases are not exhaustive, and it is your responsibility to ensure that your code works in general, not just for a few supplied test cases. We may use additional hidden test cases in grading your lab assignments.

**Run the first cell, Initialize Otter**, to import the autograder and submission exporter.

## Answer cells
Throughout the assignment, replace `...` with your answers. We use `...` as a placeholder and theses should be deleted and replaced with your answers.

Your answers must be in the cells marked `# ANSWER CELL`, including your final calculation/code. However, do not perform scratchwork in `# ANSWER CELL`. Add a new cell to perform your scratchwork. Your scratchwork will not be graded and does not need to be included in your submission unless otherwise noted.

To read the documentation on a Python function, you can type `help()` and add the function name between parentheses.

## Score Breakdown
Question| Points
--- | --
1 | 5.0
2 | 5.0
3 | 5.0
4 | 5.0
5 | 5.0
Total | 25.0

**Run the cell below**, to import the required modules.

```python
# Please run this cell, and do not modify the contents
import math
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
import cartopy.crs as ccrs
np.seterr(all='ignore');

import hashlib
def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

## Question 1: Heated Copper Bar

In this problem you will examine the heat distribution in a one-dimensional copper bar over time using the heat equation. In particular, you will create a visual representation of the temperature at different points along the bar as time progresses.

The copper bar shown in Figure 1 has a fixed length $l$ and known initial temperature distribution $f(x)$. The two ends are maintained at a constant temperature of 0&deg;C.  The heat equation describes how the temperature at any point along the bar changes over time due to the flow of heat within the material. The heat equation in one dimension is

$$ \dfrac{\partial }{\partial t}u(x,t) = \alpha \dfrac{\partial^2 }{\partial x^2}u(x,t), $$

where $u(x,t)$ is the temperature at position $x \in (0,l)$ at time $t > 0$ and $\alpha$ is the thermal diffusivity (a known constant).


<img src='resources/rod.png' width=500/>

<center>Figure 1: 1D bar with homogeneous boundary conditions.</center>

The bar the following physical properties:
* length $l = 100\,\mathrm{cm}$
* constant mass density $\rho = 8.92 \,\mathrm{g/cm^3}$
* specific heat capacity $c  = 0.092 \,\mathrm{cal/( g\,^\circ C)}$
* thermal conductivity $k = 0.95 \, \mathrm{cal/(cm\,sec\,^{\circ}C})$

The thermal diffusivity is related to these physical properties through the equation $$\alpha = \dfrac{k}{c \rho}$$
Finally, assume the bar is heated to an initial temperature of
$$ u(x,0) = f(x) = 100 \sin\left(\frac{3\pi x}{l}\right),  \label{eq:u0}\tag{1} $$
where $x$ is the horizontal coordinate along the bar. 
The temperature distribution $u(x,t)$ at point $x$ and time $t$ under these conditions is then 

$$ u(x,t) = u(x,0) \exp\left(-\lambda^2 t\right),  \label{eq:uxt}\tag{2} $$
where we've defined
$$\lambda^2 = \alpha\left(\dfrac{3\pi}{l}\right)^2$$

### Question 1.0

Create a function `u0(x)` that takes in:
* `x`: `float` or `np.ndarray`
     * coordinates along the bar, $x \in (0,l)$

and returns the intital temperature distribution shown in Eq.$\eqref{eq:u0}$ at each point in `x` for this particular problem.

Create a function `u(x,t)` that takes in:
* `x`: `float` or `np.ndarray` 
    * coordinates along the bar, $x \in (0,l)$ 
* `t`:  `float` 
    * time,  $ t > 0$ 
    

and returns the temperature distribution at time `t` shown in Eq.$\eqref{eq:uxt}$ at each point in `x` for this particular problem.

```python
# ANSWER CELL

# Bar length
L = ...

# Density gm/cm3
rho = ...

# Specific heat
c = ...

# Thermal conductivity
k = ...

# Initial temperature distribution function u0(x)

# Temperature distribution function  u(x,t)
```

```python
grader.check("q1.0")
```

### Question 1.1

We can now visualize the temperature distribution along the bar as time progresses. To this end, first create a one-dimensional grid `x` of 1000 equally spaced points in the range $\left[0,l\right]$.
Then, create a figure stored in `fig_1` and perform the following tasks on the same figure:
* Plot the temperature distribution at time $t = 0\,\mathrm{s}$ with a blue (`b`) solid (`-`) line . 
* Plot the temperature distribution at time $t = 50\,\mathrm{s}$ with a red (`r`) dashed (`--`) line. 
* Plot the temperature distribution at time $t = 150\,\mathrm{s}$ with a green (`g`) dash-dotted (`-.`) line. 
* Set the figure title to  "Question 1" with font size 16.
* Set the x-axis label to "x, cm" with font size 14. Note the spacing.
* Set the y-axis label to "Temperature, degrees Celsius" with font size 14. Note the spacing.
* Set the x-axis limits to (0,100).
* Set the y-axis limits to (-120,120).
* Add a legend to the figure. The labels for the three line plots should respectively be "t = 0s", "t = 50s" and "t = 100s". Note the spacing.
* Add grid lines to the figure.

When you are finished, your figure should look like Figure 2. Feel free to experiment with plotting options that have not been explicitly specified.

<img src='resources/q1.jpeg' width = 700/>


<center>Figure 2: Question 1 plot.</center>

```python
# Do not modify this line for grading purposes
import matplotlib.pyplot as plt

# ANSWER CELL

# Create figure 
fig_1 = ...

# x grid 
x = ...

# Plot u(x,t) vs. x for different t values

# Add figure title

# Add axes labels

# Specify axis limits

# Add legend

# Add gridlines

plt.show()
```

```python
# Do not modify this cell for grading purposes
x_q1 = np.copy(x)
```

```python
grader.check("q1.1")
```

## Question 2: Vibrating Membrane 
In this problem you will examine the natural modes of vibration of a rectangular membrane of length $a$ and width $b$, fixed along all four edges. Vibrating membranes are commonly encountered in various fields such as structural engineering, mechanical engineering and acoustics. Understanding their mode shapes is crucial for analyzing their behavior and resonance frequencies.

Assuming the membrane obeys the two-dimenional wave equation, the amplitude $f$ of the $(m,n)$ natural mode of vibration is

$$f\left(x,y,a,b\,; m,n\right) = \sin\left(\frac{m \pi}{a}x\right)\sin\left(\frac{n \pi}{b}y\right),\label{eq:modeshape}\tag{3}$$

where $x$ and $y$ are coordinates along the memebrane and $m$ and $n$ are positive integer mode indices.

### Question 2.0

Write a function `mode_shape(x,y,a,b,m,n)` which takes in:
* `x`: `float` or `np.ndarray`
     * $x$ coordinates along the membrane, $x \in (0,a)$
* `y`: `float` or `np.ndarray`
     * $y$ coordinates along the membrane, $y \in (0,b)$
* `a`: `float`
     * length of membrane, $a > 0$
* `b`: `float`
     * width of membrane, $b > 0$
* `m`: `int`
     * mode index along x, $m = 1,2,3,\dots$
* `n`: `int`
     * mode index along y, $n = 1,2,3,\dots$

and returns the mode amplitude $f\left(x,y,a,b\,; m,n\right)$ shown in Eq.$\eqref{eq:modeshape}$ at each point defined by `x` and `y`.

```python
# ANSWER CELL
```

```python
grader.check("q2.0")
```

### Question 2.1

We are now ready to visualize the membrane's natural modes of vibration. To this end, first create a two-dimensional grid of 1000 equally spaced points in each direction. The documentation of `meshgrid` and `plt.plot_surface` may be helpful.

Then, create a figure stored in `fig_2` and perform the following tasks on the same figure:
* Make a surface plot of the natural mode shape for $a = m = 3$ and $b= n = 2$. 
* Set the surface colormap to `plasma`.
* Set the figure title to  "Natural mode shape for m = 3, n = 2" with font size 14.
* Set the x-axis label to "X" with font size 12.
* Set the y-axis label to "Y" with font size 12.
* Set the y-axis label to "Z" with font size 12.
* Adjust the plot to have equal axis aspect ratios.


When you are finished, your figure should look like Figure 2. Feel free to experiment with plotting options that have not been explicitly specified.

<img src="resources/q2_1.png" width=450/>

<center>Figure 3: Question 2.1 plot.</center>

```python
# ANSWER CELL
# Do not modify this line for grading purposes
import matplotlib.pyplot as plt

# Length
a = ...
# Width
b = ...
# Parameter m
m = ...
# Parameter n
n = ...

# Create x- and y-ranges
x = ...
y = ...

# Create combined x and y grid
X,Y = ...

# Create plot
...
```

```python
# Do not modify this cell for grading purposes
x_q2_1 = np.copy(x)
y_q2_1 = np.copy(y)
X_q2_1 = np.copy(X)
Y_q2_1 = np.copy(Y)
```

```python
grader.check("q2.1")
```

### Question 2.2


Create a two-dimensional grid of 1000 equally spaced $x$ and $y$ grid coordinates for a square memebrane of unit length $a = b = 1$. Then, create a figure stored in `fig_3` and perform the following tasks on the same figure:
* Create a 2x2 subplot layout and plot the membrane mode shape for $m ,n = {1,2}$. Store the subplot axes be in the variable `axs`. The subplot in the $m$-th row and $n$-th column should contain the mode shape for the respective $m$ and $n$ values. For example, the subplot in the lower left (row 2, column 1) should contain the plot for $m = 2$ and $n =1$.
* Set the figure title (suptitle) title to "Mode shapes" with font size 16.
* For each plot do the following:
    * Set the surface colormap to `coolwarm`
    * Set the subplot title "m = {m}, n = {n}" where you insert the appropriate value in the braces. See Figure 4 for reference.
    * Set the x-axis label to "X".
    * Set the y-axis label to "Y".
    * Set the y-axis label to "Z".


When you are finished, your figure should look like Figure 4. Feel free to experiment with plotting options that have not been explicitly specified.


<img src="resources/q2_2.png" width = 600/>


<center>Figure 4: Question 2.2 plot.</center>

```python
# ANSWER CELL

# Do not modify this line for grading purposes
import matplotlib.pyplot as plt

# Create x- and y-ranges
x = ...
y = ...

# Create combined x and y grid
X,Y = ...

# Create 3D figure subplots
fig_3, axs = ...

#

# Add figure title
...

plt.show()
```

```python
# Do not modify this cell for grading purposes
x_q2_2 = np.copy(x)
y_q2_2 = np.copy(x)
X_q2_2 = np.copy(X)
Y_q2_2 = np.copy(Y)
```

```python
grader.check("q2.2")
```

## Question 3: Potential flow around a circular cylinder
 
Potential flow around a circular cylinder is a classical problem in fluid dynamics. It inovles the flow of an inviscid, incompressible fluid around a circular cylinder transverse to the flow. Far from the cylinder, the flow is unidirectional and uniform with a constant velocity $U$ in the x-direction. The flow has no vorticity and thus the velocity field is irrotational which allows for closed-form analytical solutions for quantities of interest.

<img src="resources/flow.gif" width=500/>

<center>Figure 5: Potential flow around a circular cylinder.</center>

### Question 3.0

First, we'll determine the pressure field. Let the coordinate system be centered on the cicular cylinder with radius $R$. The solution for the (gauge) *pressure* field $p$ in polar coordinates is then 
$$ p = \frac{1}{2} \rho U^2 \left( 2 \frac{R^2}{r^2}\cos\left(2\theta\right) - \frac{R^4}{r^4}\right) \label{eq:pressure}\tag{4} $$ 
Another quantity of interest is the so-called *stream function* $\psi$. It can be used to plot streamlines, which represent the trajectories of particles in a steady flow. The stream function for this problem is
$$\psi  = U\left(r-{\frac {R^{2}}{r}}\right)\sin \theta \label{eq:stream}\tag{5} $$
Finally, we are also interested in the *velocity potential* $\phi$. We'll use it to plot equipotential lines along which the fluid has uniform velocity. The velocity potential for this problem is
$$ \phi = Ur\left(1+{\frac {R^{2}}{r^{2}}}\right)\cos \theta  \label{eq:velpotential}\tag{6} $$

Let the cylinder radius have radius $R = 1$ and the far field fluid velocity be $U = 1000$. You are tasked with computing the three quantities described above on a uniform two-dimensional **polar** grid. The polar coordinates are the radius $r$ and angle $\theta$. The grid should have 1000 equally spaced points in the radial direction within the range $\left[R, 10R\right]$ and 1000 equally spaced points in the angular direction within the range $\left[0, 2\pi\right]$.

```python
# Fluid velocity away from the cylinder
U = ...

# Cylinder radius
R = ...

# Create r and theta 1D grids
theta = ...
rr = ...

# Create meshgrid in two dimensions
...

# Compute pressure around the cylinder, Equation 4
pressure = ...

# Compute stream function, Equation 5
stream = ...

# Compute velocity potential , Equation 6
velocity_potential = ...
```

```python
grader.check("q3.0")
```

### Question 3.1

Now, we will plot the pressure, streamlines and velocty equipotential lines. First, we'll need to convert the 2D polar grid to Cartesian coordinates. The relatonship between polar and Cartesian coordiantes is
\begin{align}
x &= r \cos{\theta} \\
y &= r \sin{\theta}
\end{align}

After you convert the grid, create a figure stored in `fig_4` and perform the following tasks on the same figure:
* Create a filled contour plot (`contourf`) of the pressure field and assign it to the variable `pressure_contour`. Set the number of levels to 50 and the colormap to `jet`.
* Add a colorbar of the pressure field plot to the figure and assign it to the variable `cbar`. The colorbar should have only two ticks - at the minimum and maximum value which occur in the plot and the labels should respectively be "Low" and "High".
* Create a contour plot (`contour`) of the streamlines and assign it to the variable `stream_contour`.. Set the number of levels to 50, the line width to 0.8 and the line color to black (`k`).
* Create a contour plot (`contour`) of the velocity potential and assign it to the variable `velocity_contour`. Set the number of levels to 50, the line width to 0.8 and the line color to white (`w`).
* Set the title to "Pressure field" with font size 14.
* Set the x-axis label to "x" with font size 12.
* Set the y-axis label to "y" with font size 12.
* Set the x-axis limits to (-5,5).
* Set the y-axis limits to (-5,5).
* Remove the ticks from the x and y axes.

*Hint:* The `negative_linestyle` may be useful for the `contour` plots.


When you are finished, your figure should look like Figure 4. Feel free to experiment with plotting options that have not been explicitly specified.


<img src="resources/q3_1.png" width = 600/>


<center>Figure 6: Question 3.1 plot.</center>

```python
# ANSWER CELL

# Do not modify this line for grading purposes
import matplotlib.pyplot as plt

# Convert polar to cartesian coordinates
X, Y = ...

# Create empty figure
fig_4 = ...

# Plot circle/cylinder, do not modify
th =  np.linspace(0,2*np.pi,100)  
xunit = np.cos(th)
yunit = np.sin(th)
h = plt.plot(xunit, yunit,'-k')

# Pressure plot
pressure_contour = ...

# Create colorbar
cbar = ...
...

# Streamline contour plot
stream_contour = ...
# Velocity equipotential lines contour plot
velocity_contour = ...

# Set title
...
# Set axes labels
...
...
# Set axes ticks
...
...
# Set axes limits 
...
...
plt.show()
```

```python
# Do not modify this cell for grading purposes
X_q3_1 = np.copy(X)
Y_q1_1 = np.copy(Y)
cbar_ticks_q3_1 = cbar.get_ticks()
cbar_ax_q3_1 = cbar.ax
```

```python
grader.check("q3.1")
```

### Question 3.2

Finally, we will plot the fluid velocity magnitude and vector field around the cylinder. The velocty vector components in polar coordiantes are

$$v_{r}={\frac {\partial \phi }{\partial r}}=U\left(1-{\frac {R^{2}}{r^{2}}}\right)\cos \theta, \label{eq:vr}\tag{7}$$

$$v_{\theta }={\frac {1}{r}}{\frac {\partial \phi }{\partial \theta }}=-U\left(1+{\frac {R^{2}}{r^{2}}}\right)\sin \theta, \label{eq:vt}\tag{8} $$

Compute these quantities on the 2D polar grid from Question 3.0 and do not modify the code which converts them to the Cartesian components $v_x$ and $v_y$. The magnitude of a 2D vector $\vec{v}$ with components $v_x$ and $v_y$ is defined as:
$$ ||\vec{v}|| = \sqrt{v_x^2 + v_y^2} , \label{eq:vnorm}\tag{9}$$

When you are done, create a figure stored in `fig_5` and perform the following tasks on the same figure:
* Create a filled contour plot (`contourf`) of the velocty magnitude $\eqref{eq:vnorm}$ and assign it to the variable `v_contour`. Set the number of levels to 30 and the colormap to `jet`.
* Add a colorbar of the velocity magniude plot to the figure and assign it to the variable `cbar`. The colorbar should have only two ticks - at the minimum and maximum value which occur in the plot and the labels should respectively be "Low" and "High".
* Create a vector plot (`quiver`) of the velocity field $v_x, v_y$ and assign it to the variable `v_quiver`. Becasue the grid is so dense, plot at only every 25th grid point in each direction (*Hint:* Use step slicing). Set the scale factor to 20000.
* Set the title to "Velocity field" with font size 14.
* Set the x-axis label to "x" with font size 12.
* Set the y-axis label to "y" with font size 12.
* Set the x-axis limits to (-3,3).
* Set the y-axis limits to (-3,3).
* Remove the ticks from the x and y axes.

*Hint:* The `negative_linestyle` may be useful for the `contour` plots.


When you are finished, your figure should look like Figure 7. Feel free to experiment with plotting options that have not been explicitly specified.


<img src="resources/q3_2.png" width = 600/>


<center>Figure 7: Question 3.2 plot.</center>



```python
# ANSWER CELL

# Do not modify this line for grading purposes
import matplotlib.pyplot as plt

# Velocity in polar coordinates
vr = ...
vt = ...

# Velocity in cartesian coordinates, do not modify these lines
vx = vr*np.cos(t) - vt*np.sin(t)
vy = vr*np.sin(t) + vt*np.cos(t)

# Create empty figure
fig_5 = ...

# Plot circle/cylinder, do not modify
th =  np.linspace(0,2*np.pi,100)  
xunit = np.cos(th)
yunit = np.sin(th)
h = plt.plot(xunit, yunit,'-k')

# Compute velocity magnitude
v_mag = ...

# Velocity magnitude filled contour plot
v_contour = ...
                  
# Add velocity magnitude colorbar
cbar = ...
...
                  
# Velocity vector field quiver plot
step = ...
v_quiver = ...

# Set title
...
# Set axes labels
...
...
# Set axes ticks
...
...
# Set axes limits 
...
...

plt.show()
```

```python
# Do not modify this cell for grading purposes
cbar_ticks_q3_2 = cbar.get_ticks()
cbar_ax_q3_2 = cbar.ax
```

```python
grader.check("q3.2")
```

## Question 4: Wind over North America

In this problem you will examine and visualize a sample dataset of air currents over North America. 

### Question 4.0

The `resources` folder contains the following data files on a 35 x 41 grid:

* `wind_x.csv`: $x$ coordinates (longitude) 
* `wind_y.csv`: $y$ coordinates (latitude) 
* `wind_u.csv`: $x$ components of velocity
* `wind_v.csv`: $y$ components of velocity

First, read the data files listed above using `pd.read_csv()` and convert them to NumPy arrays.

*Hint:* Be careful with the `header` argument.

```python
# ANSWER CELL

# Do not modify this line for grading purposes
import pandas as pd

# Read wind_x data
wind_x = ...
# Read wind_y data
wind_y = ...
# Read wind_u data
wind_u = ...
# Read wind_v data
wind_v = ...
```

```python
grader.check("q4.0")
```

### Question 4.1

Compute the magnitude of the wind velocity at each grid point and write it to `resources/wind_mag.csv` using `pd.DataFrame.to_csv()`. Make sure to preserve the shape of the wind velocity arrays.

The magnitude of a 2D vector $\vec{v}$ with components $v_x$ and $v_y$ is defined as:
$$ ||\vec{v}|| = \sqrt{v_x^2 + v_y^2}$$

*Hint:* Be careful with the `header` and `index` arguments.

```python
# ANSWER CELL

# Compute magnitude of wind velocity on the grid
wind_mag = ...

# Write the wind magnitude data to a csv file
```

```python
grader.check("q4.1")
```

### Question 4.2

Now, we'll plot the wind data overlaid on a map of North America. To this end, we will use the `cartopy` package (refer to [Chapter 12.3](https://pythonnumericalmethods.berkeley.edu/notebooks/chapter12.03-Working-with-Maps.html) of the course textbook).

After you've familiarized yourself with the `cartopy` package, create a figure stored in `fig_6` and perform the following tasks on the same figure:
* Create axes using the Plate Carrée projection and draw coastlines on the plot. 
* Create a filled contour plot (`contourf`) of the velocty magnitude and assign it to the variable `wind_contour`. Set the number of levels to 20.
* Add a colorbar of the contour plot to the figure. Set the scale (`shrink`) factor to an appropriate value.
* Add a grid to the figure with transparency (`alpha`) of 0.5 and assign it to the variable `grid`. The gird lines should also be dotted (`:`) and black (`k`). Also, draw longitude and latitude tick labels only on the bottom and right sides. Refer to the documentation and [examples](https://scitools.org.uk/cartopy/docs/latest/gallery/gridlines_and_labels/gridliner.html). 
* Set the title to "Wind velocity magnitude".
* Set the x-axis limits to the minimum and maximum values in wind_x.csv.
* Set the y-axis limits to the minimum and maximum values in wind_y.csv.
 
When you are finished, your figure should look like Figure 8. Feel free to experiment with plotting options that have not been explicitly specified.


<img src="resources/q4_2.png" width = 600/>


<center>Figure 8: Question 4.2 plot.</center>

```python
# ANSWER CELL

# Do not modify these lines for grading purposes
import matplotlib.pyplot as plt
import cartopy.crs as ccrs

# Create figure
fig_6 = ...

# Create axes using Plate Carrée projection
ax = ...

# Add coastilines
...

# Create wind magnitude filled contour plot

# Create colorbar
cbar = ...

# Set limits/extent of plot
...

# Add grid lines
grid = ...

# Set title
...

plt.show()
```

```python
grader.check("q4.2")
```

### Question 4.3

Lastly, we'll plot the wind vector field using `quiver`. Create a figure stored in `fig_7` and perform the following tasks on the same figure:
* Create axes using the Plate Carrée projection and draw coastlines on the plot.
* Create a vector plot (`quiver`) of the velocty field and assign it to the varaible `vel_quiver`. Feel free to experiment with the colormap but use the `wind_mag` array as the colormapping values of the arrows. 
* Add a colorbar of the contour plot to the figure and assign it to the varaible `cbar`. Set the title to "Magnitude" and the scale (`shrink`) factor to an appropriate value.
* Add a grid to the figure with transparency (`alpha`) of 0.5 and assign it to the variable `grid`. The gird lines should also be dotted (`:`) and black (`k`). Also, draw longitude and latitude tick labels only on the bottom and left sides. Refer to the documentation and [examples](https://scitools.org.uk/cartopy/docs/latest/gallery/gridlines_and_labels/gridliner.html). 
* Set the title to "Wind velocity field".
* Set the x-axis limits to the minimum and maximum values in wind_x.csv.
* Set the y-axis limits to the minimum and maximum values in wind_y.csv.
 
When you are finished, your figure should look like Figure 9. Feel free to experiment with plotting options that have not been explicitly specified.


<img src="resources/q4_3.png" width = 600/>


<center>Figure 9: Question 4.3 plot.</center>

Finally, plot the wind vector field. using the function `quiver`

```python
# ANSWER CELL

# Do not modify these lines for grading purposes
import matplotlib.pyplot as plt
import cartopy.crs as ccrs

# Create figure
fig_7 = ...

# Create axes using Plate Carrée projection
ax = ...

# Add coastilines
...

# Create wind magnitude filled contour plot
# Create colorbar

# Add colorbar
cbar = ...

# Set limits/extent of plot
...

# Add grid lines
grid = ...

# Set title
...

plt.show()
```

```python
grader.check("q4.3")
```

## Problem 5: Earthquakes


In this problem you will analyze and visualize earthquake data for the continental United States in 2022.
### Question 5.0


The `resources/earthquakes_2022.csv` file contains data about earthquakes with magnitude above 2.5 in the continental United States during 2022. The data includes:

* `latitude`: decimal degrees latitude
* `longitude`: decimal degrees longitude
* `mag`: the magnitude of the event

First, read the data using `pd.read_csv` and examine it. Then, create a figure stored in `fig_8` and perform the following tasks on the same figure:
* Create a histogram `hist` with 6 bins in the range (2,8) and assign the output to the variables `values, bins, bars`. Set the bin edge color to black (`k`).
* Set the title to "Continental US earthquakes in 2022".
* Set the x-axis label to "Magnitude".
* Set the y-axis label to "Frequency".
* Label the histogram bars with the count in each bin (*Hint:* Refer to the documentation of `bar_label`.).
 
When you are finished, your figure should look like Figure 10. Feel free to experiment with plotting options that have not been explicitly specified.

<img src="resources/q5_0.png" width = 500/>

<center>Figure 10: Question 5.0 plot.</center>

```python
# ANSWER CELL

# Do not modify these lines for grading purposes
import matplotlib.pyplot as plt
import pandas as pd

# Read data
df = ...

# Create figure
fig_8 = ...

# Plot histogram
values, bins, bars = ...
 
# Set title and labels
...
...
...

# Add bar labels
...

plt.show()
```

```python
grader.check("q5.0")
```

### Question 5.1

We will now zoom in on the state of California and the surrounding area. Create a figure stored in `fig_9` and perform the following tasks on the same figure:
* Create axes using the Plate Carrée projection and draw coastlines on the plot.
* Create a scatter plot (`scatter`) of the earthquake location and assign it to the variable `eq_scatter`. Use the magnitude data for the colormapping. Set the transparency (`alpha`) of the markers to 0.5. 
* Add a colorbar of the pressure field plot to the figure and assign it to the variable `cbar`. Set the scale (`shrink`) factor to an appropriate value and set the label to 'Magnitude'.
* Add a grid to the figure with transparency (`alpha`) of 0.5 and assign it to the variable `grid`. The gird lines should also be dotted (`:`) and black (`k`). Also, draw longitude and latitude tick labels only on the bottom and left sides. Refer to the documentation and [examples](https://scitools.org.uk/cartopy/docs/latest/gallery/gridlines_and_labels/gridliner.html).  Also, draw longitude and latitude tick labels only on the bottom and left sides. Refer to the [documentation](https://scitools.org.uk/cartopy/docs/latest/gallery/gridlines_and_labels/gridliner.html). 
* Set the title to 'CA area earthquakes in 2022'.
* To create a rough bounding box around California:
    * set the x-axis limits to (-126,-114).
    * set the y-axis limits to (32,42).
 
When you are finished, your figure should look like Figure 11. Feel free to experiment with plotting options that have not been explicitly specified.


<img src="resources/q5_1.png" width = 600/>


<center>Figure 11: Question 5.1 plot.</center>

```python
# ANSWER CELL

# Do not modify these lines for grading purposes
import matplotlib.pyplot as plt
import cartopy.crs as ccrs

# Create figure
fig_9 = ...

# Create axes using Plate Carrée projection
ax = ...

# Add coastilines
...

# Scatter plot of earthquakes
eq_scatter = ...

# Create colorbar
cbar = ...

# Set axes extent/limits    
...

# Add grid lines
grid = ...

# Set title
...

plt.show()
```

```python
grader.check("q5.1")
```

### You're done with this Lab!

**Important submission information:** After completing the assignment, click on the Save icon from the Tool Bar &nbsp;<i class="fa fa-save" style="font-size:16px;"></i>&nbsp;. After saving your notebook, **run the cell with** `grader.check_all()` and confirm that you pass the same tests as in the notebook. Then, **run the final cell** `grader.export()` and click the link to download the zip file. Then, go to Gradescope and submit the zip file to the corresponding assignment. 

**Once you have submitted, stay on the Gradescope page to confirm that you pass the same tests as in the notebook.**

```python
import matplotlib.image as mpimg
img = mpimg.imread('resources/animal.png')
imgplot = plt.imshow(img)
imgplot.axes.get_xaxis().set_visible(False)
imgplot.axes.get_yaxis().set_visible(False)
print("Congrats on finishing Lab 6!")
plt.show()
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1.0 (points: 1)
```
>>> # Check scalar input
>>> assert np.isclose(u0(0),0,rtol = 1e-2), 'Check your u0 function.'
>>> assert np.isclose(u0(10),80.90169943749473,rtol = 1e-2), 'Check your u0 function.'
>>> assert np.isclose(u0(60),-58.778525229247336,rtol = 1e-2), 'Check your u0 function.'
>>> assert np.isclose(u0(100),0,rtol = 1e-2), 'Check your u0 function.'
>>> 
>>> # Check array input
>>> x = np.linspace(0,100)
>>> assert len(u0(x)) == 50, 'Check your u0 function.'
>>> assert np.isclose(np.sum(u0(x)),1036.604609717437,rtol = 1e-2), 'Check your u0 function.'
```
```
>>> # Check scalar input
>>> assert np.isclose(u(0,100),0,rtol = 1e-2), 'Check your u function.'
>>> assert np.isclose(u(10,10),72.99612859396403,rtol = 1e-2), 'Check your u function.'
>>> assert np.isclose(u(60,30),-43.17628127422244,rtol = 1e-2), 'Check your u function.'
>>> assert np.isclose(u(100,100),0,rtol = 1e-2), 'Check your u function.'
>>> 
>>> # Check array input
>>> x = np.linspace(0,100)
>>> assert len(u(x,10)) == 50, 'Check your u function.'
>>> assert np.isclose(np.sum(u(x,10)),935.3094424239099,rtol = 1e-2), 'Check your u function.'
```


Tests for q1.1 (points: 4)
```
>>> # Check x grid
>>> assert len(x_q1) == 1000
>>> assert np.isclose(np.sum(x_q1),50000,rtol = 1e-2)
>>> 
>>> # Check linestyles
>>> assert fig_1.axes[0].get_lines()[0].get_linestyle() == '-'
>>> assert fig_1.axes[0].get_lines()[1].get_linestyle() == '--'
>>> assert fig_1.axes[0].get_lines()[2].get_linestyle() == '-.'
>>> 
>>> # Check line colors
>>> assert fig_1.axes[0].get_lines()[0].get_color() == 'b'
>>> assert fig_1.axes[0].get_lines()[1].get_color() == 'r'
>>> assert fig_1.axes[0].get_lines()[2].get_color() == 'g'
>>> 
>>> # Check line labels
>>> assert fig_1.axes[0].get_lines()[0].get_label() == 't = 0s'
>>> assert fig_1.axes[0].get_lines()[1].get_label() == 't = 50s'
>>> assert fig_1.axes[0].get_lines()[2].get_label() == 't = 150s'
>>> 
>>> # Check xlabel and fontsize
>>> assert fig_1.axes[0].get_xlabel() == 'x, cm', 'Check the x-axis label.'
>>> assert fig_1.axes[0].xaxis.label.get_fontsize() == 14, 'Check the x-axis label fontsize.'
>>> 
>>> # Check ylabel and fontsize
>>> assert fig_1.axes[0].get_ylabel() == 'Temperature, degrees Celsius', 'Check the y-axis label.'
>>> assert fig_1.axes[0].yaxis.label.get_fontsize() == 14, 'Check the y-axis label fontsize.'
>>> 
>>> # Check title and fontsize
>>> assert fig_1.axes[0].get_title() == 'Question 1', 'Check the title label.'
>>> assert fig_1.axes[0].title.get_fontsize() == 16, 'Check the title label fontsize.'
>>> 
>>> # Check legend is added
>>> assert type(fig_1.axes[0].get_legend()) != type(None), 'Make sure to add a legend.'
>>> 
>>> # Check xlim and ylim
>>> assert fig_1.axes[0].get_xlim() == (0.0, 100.0), 'Check the x-axis limits.'
>>> assert fig_1.axes[0].get_ylim() == (-120.0, 120.0), 'Check the y-axis limits.'
```


Tests for q2.0 (points: 1)
```
>>> # Check scalar input
>>> assert np.isclose(mode_shape(0.2,0.6,1,1,1,1),0.5590169943749475,rtol = 1e-2), 'Check your mode_shape function.'
>>> assert np.isclose(mode_shape(0.23,0.7,1,1,3,4),0.486145764016177,rtol = 1e-2), 'Check your mode_shape function.'
>>> assert np.isclose(mode_shape(0.3,0.3,1,1,8,1),0.7694208842938134,rtol = 1e-2), 'Check your mode_shape function.'
>>> assert np.isclose(mode_shape(0.7,0.8,1,1,5,5),0,rtol = 1e-3), 'Check your mode_shape function.'
>>> 
>>> # Check array input
>>> x = np.linspace(0,100)
>>> assert len(mode_shape(x,x,1,1,1,1)) == 50, 'Make sure your mode_shape function works for x and y arrays.'
>>> assert np.isclose(np.sum(mode_shape(x,x,1,1,1,1)),24.499999999999982,rtol = 1e-2), 'Make sure your mode_shape function works for x and y arrays.'
```


Tests for q2.1 (points: 2)
```
>>> # Check 1D x grid
>>> assert len(x_q2_1) == 1000, 'Check your 1D x grid.'
>>> assert np.isclose(np.sum(x_q2_1),1500,rtol = 1e-2), 'Check your 1D x grid.'
>>> 
>>> # Check 1D y grid
>>> assert len(y_q2_1) == 1000, 'Check your 1D y grid.'
>>> assert np.isclose(np.sum(y_q2_1),1000,rtol = 1e-2), 'Check your 1D y grid.'
>>> 
>>> # Check 1D x grid
>>> assert X_q2_1.shape == (1000, 1000), 'Check your 2D x grid.'
>>> assert np.isclose(np.sum(X_q2_1),1500000,rtol = 1e-2), 'Check your 2D x grid.'
>>> 
>>> # Check 1D y grid
>>> assert Y_q2_1.shape == (1000, 1000), 'Check your 2D y grid.'
>>> assert np.isclose(np.sum(Y_q2_1),1000000,rtol = 1e-2), 'Check your 2D y grid.'
```
```
>>> # Check colormap
>>> assert fig_2.axes[0].get_children()[0].get_cmap().name == 'plasma', 'Check colormap.'
>>> 
>>> # Check aspect ratio 
>>> assert fig_2.axes[0].get_aspect() == 'equal', 'Check the axes aspect ratio.'
>>> 
>>> # Check xlabel and fontsize
>>> assert fig_2.axes[0].get_xlabel() == 'X', 'Check the x-axis label.'
>>> assert fig_2.axes[0].xaxis.label.get_fontsize() == 12, 'Check the x-axis label fontsize.'
>>> 
>>> # Check ylabel and fontsize
>>> assert fig_2.axes[0].get_ylabel() == 'Y', 'Check the y-axis label.'
>>> assert fig_2.axes[0].yaxis.label.get_fontsize() == 12, 'Check the y-axis label fontsize.'
>>> 
>>> # Check zlabel and fontsize
>>> assert fig_2.axes[0].get_zlabel() == 'Z', 'Check the z-axis label.'
>>> assert fig_2.axes[0].zaxis.label.get_fontsize() == 12, 'Check the y-axis label fontsize.'
>>> 
>>> # Check title and fontsize
>>> assert fig_2.axes[0].get_title() == 'Natural mode shape for m = 3, n = 2', 'Check the title label.'
>>> assert fig_2.axes[0].title.get_fontsize() == 14, 'Check the title label fontsize.'
```


Tests for q2.2 (points: 2)
```
>>> # Check 1D x grid
>>> assert len(x_q2_2) == 1000, 'Check your 1D x grid.'
>>> assert np.isclose(np.sum(x_q2_2),500,rtol = 1e-2), 'Check your 1D x grid.'
>>> 
>>> # Check 1D y grid
>>> assert len(y_q2_2) == 1000, 'Check your 1D y grid.'
>>> assert np.isclose(np.sum(y_q2_2),500,rtol = 1e-2), 'Check your 1D y grid.'
>>> 
>>> # Check 1D x grid
>>> assert X_q2_2.shape == (1000, 1000), 'Check your 2D x grid.'
>>> assert np.isclose(np.sum(X_q2_2),500000,rtol = 1e-2), 'Check your 2D x grid.'
>>> 
>>> # Check 1D y grid
>>> assert len(Y_q2_2) == 1000, 'Check your 2D y grid.'
>>> assert np.isclose(np.sum(Y_q2_2),500000,rtol = 1e-2), 'Check your 2D y grid.'
```
```
>>> # Check title
>>> assert fig_3._suptitle.get_text() == 'Mode shapes', 'Check figure suptitle.'
>>> assert fig_3._suptitle.get_fontsize() == 16, 'Check figure suptitle font size.'
>>> # Check subplots
>>> assert axs.shape == (2, 2), 'Make sure you have a 2x2 subplot layout stored in the variable "axs".'
>>> lst = [1,2]
>>> for m in lst:
...     for n in lst: 
...         assert axs[m-1,n-1].get_xlabel() == 'X', 'Check x-axis labels.'
...         assert axs[m-1,n-1].get_ylabel() == 'Y', 'Check y-axis labels.'
...         assert axs[m-1,n-1].get_zlabel() == 'Z', 'Check z-axis labels.'
...         assert axs[m-1,n-1].get_title() == 'm = {m}, n = {n}'.format(m=m,n=n), 'Check title labels.'
...         assert axs[m-1,n-1].get_children()[0].get_cmap().name == 'coolwarm', 'Check colormap.'
```


Tests for q3.0 (points: 1)
```
>>> # Check grids 
>>> assert np.isclose(sum(theta), 3141.592653589793,rtol = 1e-2), 'Check theta 1D grid.'
>>> assert np.isclose(sum(rr), 5500.000000000001,rtol = 1e-2), 'Check r 1D grid.'
>>> assert np.isclose(sum(sum(t)), 3141592.6535897925,rtol = 1e-2), 'Check r 2D grid.'
>>> assert np.isclose(sum(sum(r)), 5500000.000000001,rtol = 1e-2), 'Check theta 2D grid.'
>>> assert np.isclose(sum(sum(pressure)), -37265.23985116976,rtol = 1e-2), 'Check pressure 2D grid.'
>>> assert np.isclose(sum(sum(stream)), 6.086685410137981e-08,rtol = 1e-2), 'Check stream 2D grid.'
>>> assert np.isclose(sum(sum(velocity_potential)), 5756137.688560322,rtol = 1e-2), 'Check velocity_potential 2D grid.'
```


Tests for q3.1 (points: 2)
```
>>> # Check 1D x grid
>>> assert X_q3_1.shape == (1000, 1000), 'Check your 2D x grid.'
>>> assert np.isclose(np.sum(X_q3_1),5499.999999999398,rtol = 1e-2), 'Check your 2D x grid.'
>>> 
>>> # Check 1D y grid
>>> assert Y_q1_1.shape == (1000, 1000), 'Check your 2D y grid.'
>>> assert np.isclose(np.sum(Y_q1_1),8.185452315956354e-12,rtol = 1e-2), 'Check your 2D y grid.'
```
```
>>> # Check colorbar
>>> assert len(cbar_ticks_q3_1) == 2, 'Check colorbar ticks.'
>>> assert np.isclose(cbar_ticks_q3_1[0],-3,rtol=1e-2) , 'Check colorbar ticks.'
>>> assert np.isclose(cbar_ticks_q3_1[1],1,rtol=1e-2) , 'Check colorbar ticks.'
>>> assert cbar_ax_q3_1.get_yticklabels()[0].get_text() == 'Low', 'Check colorbar ticks.'
>>> assert cbar_ax_q3_1.get_yticklabels()[1].get_text() == 'High', 'Check colorbar ticks.'
>>> 
>>> # Check pressure plot
>>> assert pressure_contour.get_cmap().name == 'jet', 'Check pressure_contour colormap.'
>>> assert 45 < len(pressure_contour.levels) < 55, 'Check pressure_contour levels.'
>>> 
>>> # Check stream plot
>>> assert stream_contour.linewidths == 0.8, 'Check stream_contour line width.'
>>> assert 45 < len(stream_contour.levels) < 55, 'Check stream_contour levels.'
>>> assert stream_contour.tcolors[0][0] == (0,0,0,1), 'Check stream_contour line colors.'
>>> 
>>> # Check velocity plot
>>> assert velocity_contour.linewidths == 0.8, 'Check velocity_contour line width.'
>>> assert 45 < len(velocity_contour.levels) < 55, 'Check velocity contour levels.'
>>> assert velocity_contour.tcolors[0][0] == (1,1,1,1), 'Check stream_contour line colors.'
```
```
>>> # Check xlabel and fontsize
>>> assert fig_4.axes[0].get_xlabel() == 'x', 'Check the x-axis label.'
>>> assert fig_4.axes[0].xaxis.label.get_fontsize() == 12, 'Check the x-axis label fontsize.'
>>> 
>>> # Check ylabel and fontsize
>>> assert fig_4.axes[0].get_ylabel() == 'y', 'Check the y-axis label.'
>>> assert fig_4.axes[0].yaxis.label.get_fontsize() == 12, 'Check the y-axis label fontsize.'
>>> 
>>> # Check title and fontsize
>>> assert fig_4.axes[0].get_title() == 'Pressure field', 'Check the title label.'
>>> assert fig_4.axes[0].title.get_fontsize() == 14, 'Check the title label fontsize.'
>>> 
>>> # Check xlim and ylim
>>> assert fig_4.axes[0].get_xlim() == (-5, 5), 'Check the x-axis limits.'
>>> assert fig_4.axes[0].get_ylim() == (-5, 5), 'Check the y-axis limits.'
>>> 
>>> # Check xticks and yticks
>>> assert len(fig_4.axes[0].get_xticks()) == 0, 'Check the x-axis ticks.'
>>> assert len(fig_4.axes[0].get_yticks()) == 0, 'Check the y-axis ticks.'
```


Tests for q3.2 (points: 2)
```
>>> # Check vr
>>> assert vr.shape == (1000, 1000), 'Check your 2D vr grid.'
>>> assert np.isclose(np.sum(vr),899593.5000242745,rtol = 1e-2), 'Check your 2D vr grid.'
>>> 
>>> # Check vt
>>> assert vt.shape == (1000, 1000), 'Check your 2D vt grid.'
>>> assert np.isclose(np.sum(vt),-1.3387762010097504e-09,rtol = 1e-2), 'Check your 2D vt grid.'
```
```
>>> # Check quiver plot
>>> assert v_quiver.scale == 20000, 'Checl quiver plot scale.'
>>> assert len(v_quiver.X) == 1600, 'Make sure to plot at every 25th grid point only.'
>>> assert np.isclose(np.sum(v_quiver.X), 8.611321685231154, rtol = 1e-2), 'Make sure to plot at every 25th grid point only.'
>>> assert len(v_quiver.Y) == 1600, 'Make sure to plot at every 25th grid point only.'
>>> assert np.isclose(np.sum(v_quiver.Y), -0.6511651628725481, rtol = 1e-2), 'Make sure to plot at every 25th grid point only.'
```
```
>>> # Check colorbar
>>> assert len(cbar_ticks_q3_2) == 2, 'Check colorbar ticks.'
>>> assert np.isclose(cbar_ticks_q3_2[0],0,rtol=1e-2) , 'Check colorbar ticks.'
>>> assert np.isclose(cbar_ticks_q3_2[1],1999.99,rtol=1e-2) , 'Check colorbar ticks.'
>>> assert cbar_ax_q3_2.get_yticklabels()[0].get_text() == 'Low', 'Check colorbar ticks.'
>>> assert cbar_ax_q3_2.get_yticklabels()[1].get_text() == 'High', 'Check colorbar ticks.'
>>> 
>>> # Check pressure plot
>>> assert v_contour.get_cmap().name == 'jet', 'Check pressure_contour colormap.'
>>> assert 25 < len(v_contour.levels) < 35, 'Check pressure_contour levels.'
>>> 
>>> # Check stream plot
>>> assert stream_contour.linewidths == 0.8, 'Check stream_contour line width.'
>>> assert 45 < len(stream_contour.levels) < 55, 'Check stream_contour levels.'
>>> assert stream_contour.tcolors[0][0] == (0,0,0,1), 'Check stream_contour line colors.'
>>> 
>>> # Check velocity plot
>>> assert velocity_contour.linewidths == 0.8, 'Check velocity_contour line width.'
>>> assert 45 < len(velocity_contour.levels) < 55, 'Check velocitycontour levels.'
>>> assert velocity_contour.tcolors[0][0] == (1,1,1,1), 'Check stream_contour line colors.'
```
```
>>> # Check xlabel and fontsize
>>> assert fig_5.axes[0].get_xlabel() == 'x', 'Check the x-axis label.'
>>> assert fig_5.axes[0].xaxis.label.get_fontsize() == 12, 'Check the x-axis label fontsize.'
>>> 
>>> # Check ylabel and fontsize
>>> assert fig_5.axes[0].get_ylabel() == 'y', 'Check the y-axis label.'
>>> assert fig_5.axes[0].yaxis.label.get_fontsize() == 12, 'Check the y-axis label fontsize.'
>>> 
>>> # Check title and fontsize
>>> assert fig_5.axes[0].get_title() == 'Velocity field', 'Check the title label.'
>>> assert fig_5.axes[0].title.get_fontsize() == 14, 'Check the title label fontsize.'
>>> 
>>> # Check xlim and ylim
>>> assert fig_5.axes[0].get_xlim() == (-3, 3), 'Check the x-axis limits.'
>>> assert fig_5.axes[0].get_ylim() == (-3, 3), 'Check the y-axis limits.'
>>> 
>>> # Check xticks and yticks
>>> assert len(fig_5.axes[0].get_xticks()) == 0, 'Check the x-axis ticks.'
>>> assert len(fig_5.axes[0].get_yticks()) == 0, 'Check the y-axis ticks.'
```


Tests for q4.0 (points: 1)
```
>>> # Check read data type
>>> assert isinstance(wind_x, np.ndarray), 'Make sure to convert the wind_x data to a numpy array.'
>>> assert isinstance(wind_y, np.ndarray), 'Make sure to convert the wind_y data to a numpy array.'
>>> assert isinstance(wind_u, np.ndarray), 'Make sure to convert the wind_u data to a numpy array.'
>>> assert isinstance(wind_v, np.ndarray), 'Make sure to convert the wind_v data to a numpy array.'
>>> 
>>> # Check data correctness
>>> assert int(sum(sum(wind_x))) == -146719, 'Make sure the wind_x is stored in X.'
>>> assert int(sum(sum(wind_y))) == 55606, 'Make sure the wind_y is stored in Y.'
>>> assert int(sum(sum(wind_u))) == 25314, 'Make sure the wind_u is stored in U.'
>>> assert int(sum(sum(wind_v))) == 607, 'Make sure the wind_v is stored in V.'
```


Tests for q4.1 (points: 1)
```
>>> import pandas as pd
>>> q4_1 = pd.read_csv('resources/wind_mag.csv',header=None).sum().sum()
>>> assert np.isclose(sum(sum(wind_mag)),27813.329567459838,rtol = 1e-2), 'Check the wind magnitude computation.' 
>>> assert np.isclose(q4_1,27813.329567459838,rtol = 1e-2), 'Check the write to csv command.'
```


Tests for q4.2 (points: 2)
```
>>> # Check velocity plot
>>> assert wind_contour.get_cmap().name == 'coolwarm', 'Check wind_contour colormap.'
>>> assert 15 < len(wind_contour.levels) < 25, 'Check velocitycontour levels.'
>>> 
>>> # Check title
>>> assert fig_6.axes[0].get_title() == 'Wind velocity magnitude', 'Check the title label.'
>>> 
>>> # Check grid
>>> assert grid.top_labels == False, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.bottom_labels == True, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.left_labels == True, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.right_labels == False, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.collection_kwargs['alpha'] == 0.5 , 'Check grid alpha'
>>> assert grid.collection_kwargs['linestyle'] == ':' , 'Check grid linestyle'
>>> assert grid.collection_kwargs['color'] == 'k' , 'Check grid color'
>>> 
>>> # Check xlim and ylim
>>> assert fig_6.axes[0].get_xlim() == (-134.3, -70.188), 'Check the x-axis limits.'
>>> assert fig_6.axes[0].get_ylim() == (17.5, 60.0), 'Check the y-axis limits.'
```


Tests for q4.3 (points: 1)
```
>>> # Check title
>>> assert fig_7.axes[0].get_title() == 'Wind velocity field', 'Check the title label.'
>>> 
>>> # Check colorbar 
>>> assert cbar.ax.get_ylabel() == 'Magnitude', 'Check the colorbar label.'
>>> 
>>> # Check grid
>>> assert grid.top_labels == False, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.bottom_labels == True, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.left_labels == True, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.right_labels == False, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.collection_kwargs['alpha'] == 0.5 , 'Check grid alpha'
>>> assert grid.collection_kwargs['linestyle'] == ':' , 'Check grid linestyle'
>>> assert grid.collection_kwargs['color'] == 'k' , 'Check grid color'
>>> 
>>> # Check xlim and ylim
>>> assert fig_7.axes[0].get_xlim() == (-134.3, -70.188), 'Check the x-axis limits.'
>>> assert fig_7.axes[0].get_ylim() == (17.5, 60.0), 'Check the y-axis limits.'
```
```
>>> # Check the color of vector arrows
>>> import matplotlib
>>> assert np.isclose(np.sum(matplotlib.cm.ScalarMappable(norm=vel_quiver.norm, cmap=vel_quiver.cmap).to_rgba(wind_mag).flatten()),np.sum(vel_quiver.get_facecolor().flatten()),rtol = 1e-2), 'Use wind_mag as the colormap values.'
```


Tests for q5.0 (points: 2)
```
>>> # Check histogram
>>> assert len(values) == 6, 'Check bin count.'
>>> assert len(bins) == 7, 'Check bin count.'
>>> assert sum(values) == 2634, 'Check histogram data.'
>>> assert sum(bins) == 35, 'Check bin range.'
```
```
>>> # Check xlabel 
>>> assert fig_8.axes[0].get_xlabel() == 'Magnitude', 'Check the x-axis label.'
>>> # Check ylabel 
>>> assert fig_8.axes[0].get_ylabel() == 'Frequency', 'Check the y-axis label.'
>>> # Check title 
>>> assert fig_8.axes[0].get_title() == 'Continental US earthquakes in 2022', 'Check the title label.'
```


Tests for q5.1 (points: 3)
```
>>> # Check plot data
>>> assert np.isclose(np.sum(eq_scatter.get_offsets()[:,0]),df['longitude'].sum(),rtol = 1e-2), 'Make sure you are plotting the correct data.'
>>> assert np.isclose(np.sum(eq_scatter.get_offsets()[:,1]),df['latitude'].sum(),rtol = 1e-2), 'Make sure you are plotting the correct data.'
>>> # Check the marker color
>>> import matplotlib
>>> assert np.isclose(np.sum(matplotlib.cm.ScalarMappable(norm=eq_scatter.norm, cmap=eq_scatter.cmap).to_rgba(df['mag'],alpha =0.5).flatten()),np.sum(eq_scatter.get_facecolor().flatten()),rtol = 1e-2), 'Use the earthquake magnitude as the colormap values. Make sure you set the color alpha as well.'
```
```
>>> # Check title
>>> assert fig_9.axes[0].get_title() == 'CA area earthquakes in 2022', 'Check the title label.'
>>> 
>>> # Check colorbar 
>>> assert cbar.ax.get_ylabel() == 'Magnitude', 'Check the colorbar label.'
>>> 
>>> # Check grid
>>> assert grid.top_labels == False, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.bottom_labels == True, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.left_labels == True, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.right_labels == False, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.collection_kwargs['alpha'] == 0.5 , 'Check grid alpha'
>>> assert grid.collection_kwargs['linestyle'] == ':' , 'Check grid linestyle'
>>> assert grid.collection_kwargs['color'] == 'k' , 'Check grid color'
>>> 
>>> # Check xlim and ylim
>>> assert np.isclose(np.sum(fig_9.axes[0].get_xlim()),np.sum([-126,-114]), rtol = 1e-2), 'Check the x-axis limits.'
>>> assert np.isclose(np.sum(fig_9.axes[0].get_ylim()),np.sum([32,42]), rtol = 1e-2), 'Check the y-axis limits.'
```
